Local, calcium- and reward-based synaptic learning rule that enhances dendritic nonlinearities can solve the nonlinear feature binding problem

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Public review):

Summary:

This computational modeling study builds on multiple previous lines of experimental and theoretical research to investigate how a single neuron can solve a nonlinear pattern classification task. The authors construct a detailed biophysical and morphological model of a single striatal medium spiny neuron, and endow excitatory and inhibitory synapses with dynamic synaptic plasticity mechanisms that are sensitive to (1) the presence or absence of a dopamine reward signal, and (2) spatiotemporal coincidence of synaptic activity in single dendritic branches. The latter coincidence is detected by voltage-dependent NMDA-type glutamate receptors, which can generate a type of dendritic spike referred to as a "plateau potential." In the absence of inhibitory plasticity, the proposed mechanisms result in good performance on a nonlinear classification task when specific input features are segregated and clustered onto individual branches, but reduced performance when input features are randomly distributed across branches. Interestingly, adding inhibitory plasticity improves classification performance even when input features are randomly distributed.

Strengths:

The integrative aspect of this study is its major strength. It is challenging to relate low-level details such as electrical spine compartmentalization, extrasynaptic neurotransmitter concentrations, dendritic nonlinearities, spatial clustering of correlated inputs, and plasticity of excitatory and inhibitory synapses to high-level computations such as nonlinear feature classification. Due to high simulation costs, it is rare to see highly biophysical and morphological models used for learning studies that require repeated stimulus presentations over the course of a training procedure. The study aspires to prove the principle that experimentally-supported biological mechanisms can explain complex learning.

Weaknesses:

The high level of complexity of each component of the model makes it difficult to gain an intuition for which aspects of the model are essential for its performance, or responsible for its poor performance under certain conditions. Stripping down some of the biophysical detail and comparing it to a simpler model may help better understand each component in isolation.

We greatly appreciate your recognition of the study’s integrative scope and the challenges of linking detailed biophysics to high-level computation. We acknowledge that the model’s complexity can obscure the contribution of individual components. However, as stated in the introduction the principles already have been shown in simplified theoretical models for instance in Tran-Van-Minh et al. 2015. Our aim here was to extend those ideas into a more biologically detailed setting to test whether the same principles still hold under realistic constraints. While simplification can aid intuition, we believe that demonstrating these effects in a biophysically grounded model strengthens the overall conclusion. We agree that further comparisons with reduced models would be valuable for isolating the contribution of specific components and plan to explore that in future work.

Reviewer #2 (Public review):

Summary:

The study explores how single striatal projection neurons (SPNs) utilize dendritic nonlinearities to solve complex integration tasks. It introduces a calcium-based synaptic learning rule that incorporates local calcium dynamics and dopaminergic signals, along with metaplasticity to ensure stability for synaptic weights. Results show SPNs can solve the nonlinear feature binding problem and enhance computational efficiency through inhibitory plasticity in dendrites, emphasizing the significant computational potential of individual neurons. In summary, the study provides a more biologically plausible solution to single-neuron learning and gives further mechanical insights into complex computations at the single-neuron level.

Strengths:

The paper introduces a novel learning rule for training a single multicompartmental neuron model to perform nonlinear feature binding tasks (NFBP), highlighting two main strengths: the learning rule is local, calcium-based, and requires only sparse reward signals, making it highly biologically plausible, and it applies to detailed neuron models that effectively preserve dendritic nonlinearities, contrasting with many previous studies that use simplified models.

Thank you for highlighting the biological plausibility of our calcium- and dopamine-dependent learning rule and its ability to exploit dendritic nonlinearities. Your positive assessment reinforces our commitment to refining the rule and exploring its implications in larger, more diverse settings.

Reviewer #1 (Recommendations for the authors):

Major recommendations:

P9: When introducing the excitatory learning rule, the reader is referred to the Methods. I suggest moving Figure 7A-D, "Excitatory plasticity" to be more prominently presented in the main body of the paper where the reader needs to understand it. There are errors in the current Figure 7, and wrong/confusing acronyms. The abbreviations "LTP-K" and "MP-K" are not intuitive. In A, I would spell out "LTP kernel" and "Theta_LTP adaptation". In B, I would spell out "LTD kernel" and "Theta_LTD adaptation".

We have clarified the terminology in Figure 7 by replacing “LTP-K” with “LTP kernel” and “MP-K” with “metaplasticity kernel”. While we kept Figure 7 in the Methods section to maintain the flow of the main text, we agree that an earlier introduction of the learning rule improves clarity. To that end, we added a simplified schematic to Figure 3 in the Results section, which provides readers with an accessible overview of the excitatory plasticity mechanism at the point where it is first introduced.

In C, for simplicity and clarity, I would only show the initial and updated LTP kernel and Calcium and remove the Theta_LTP adaptation curve, it's too busy and not necessary. Similarly in D, I would show only the initial and updated LTD kernel and Calcium and remove the Theta_LTD adaptation curve. In the current version of the Figure, panel B, right incorrectly labels "Theta_LTD" as "Theta_LTP". Panel D incorrectly labels "LTD kernel" as "LTP/MP-K" in the subheading and "MP/LTP-K" in the graph.

To avoid confusion and better illustrate the interactions between calcium signals, kernels, and thresholds, we have added a movie showing how these components evolve during learning. The figure panels remain as originally designed, since the LTP kernel governs both potentiation and depression through metaplastic threshold adaptation, while the LTD kernel remains fixed.

P17: Again, instead of pointing the reader to the Methods, I would move Figure 7E, "Inhibitory plasticity" to the main body of the paper where the reader needs to understand it. For clarity, I would label "C_TL" and "Theta_Inh,low" and "C_TH" as "Theta_Inh,high". The right panel could be better labeled "Inhibitory plasticity kernel". The left panel could be better labeled "Theta_Inh adaptation", with again replacing the acronyms "C_TL" and "C_TH". The same applies to Fig. 5D on P19.

We have updated the labeling in Figures 5D and 7E for clarity, including replacing "C_TL" and "C_TH" with "Theta_Inh,low" and "Theta_Inh,high". In addition, we added a simplified schematic of the inhibitory plasticity rule to Figure 5 to assist the reader’s understanding when presenting the results. Figure 7E remains in the Methods section to preserve the flow of the main text.

P12: I would suggest simplifying Fig. 3 panels and acronyms as well. Remove "MP-K" from C and D. Relabel "LTP-K" as "LTP kernel". The same applies to Fig. 5E on P19 and Fig. 3 - supplement 1 on P46 and Fig 6 - supplement 1 on P49.

We have simplified the labeling across all relevant figures by replacing “MP-K” with “metaplasticity kernel” and “LTP-K” with “LTP kernel.” To maintain clarity, we retained these terms in only one panel as a reference.

Minor recommendations:

P4: "Although not discussed much in more theoretical work, our study demonstrates the necessity of metaplasticity for achieving stable and physiologically realistic synaptic weights." This sentence is jarring. BCM and metaplasticity has been discussed in hundreds of theory papers! Cite some. This sentence would more accurately read, "Our study corroborates prior theory work (citations) demonstrating that metaplasticity helps to achieve stable and physiologically realistic synaptic weights."

We have followed the reviewers suggestion and updated the sentence to: Previous theoretical studies (Bienenstock et al., 1982; Fusi et al., 2005; Clopath et al., 2010; Benna & Fusi, 2016; Zenke & Gerstner, 2017) demonstrate the essential role of metaplasticity in maintaining stability in synaptic weight distributions. (page 2 line 49-51, page 3 line 1)

P9: Grammar. "The neuron model was during training activated..." should read "During training, the neuron model was activated..."

Corrected

P17: Lovett-Barron et al., 2012 is appropriately cited here. Milstein et al., Neuron, 2015 also showed dendritic inhibition regulates plateau potentials in CA1 pyramidal cells in vitro, and Grienberger et al., Nat. Neurosci., 2017 showed it in vivo.

P19 vs P16 vs P21. Fig. 4B, Fig. 5B, and Fig. 6B choose different strategies to show variance across seeds. Please choose one strategy and apply to all comparable plots.

We thank the reviewer for these helpful points.

We have added the suggested citations (Milstein et al., 2015; Grienberger et al., 2017) alongside Lovett-Barron et al., 2012.

Variance across seeds is now displayed uniformly (mean is solid line STD is shaded area) in Figures 4B, 5B, and 6B.

Reviewer #2 (Recommendations for the authors):

Major Points:

(1) Quality of Scientific Writing:

i. Mathematical and Implementation Details:

I appreciate the authors' efforts in clarifying the mathematical details and providing pseudocode for the learning rule, significantly improving readability and reproducibility. The reference to existing models via GitHub and ModelDB repositories is acceptable. However, I suggest enhancing the presentation quality of equations within the Methods section-currently, they are low-resolution images. Please consider rewriting these equations using LaTeX or replacing them with high-resolution images to further improve clarity.

We appreciate the reviewer’s comment regarding clarity and reproducibility. In response, we have rewritten all equations in LaTeX to improve their readability and presentation quality in the Methods section.

ii. Figure quality.

I acknowledge the authors' effort to improve figure clarity and consistency throughout the manuscript. However, I notice that the x-axis label "[Ca]_v (μm)" in Fig. 7E still appears compressed and unclear. Additionally, given the complexity and abundance of hyperparameters or artificial settings involved in your experimental design and learning rule (such as kernel parameters, metaplasticity kernels, and unspecific features), the current arrangement of subfigures (particularly Fig. 3C, D and Fig. 5D, E) still poses readability challenges. I recommend reordering subfigures to present primary results (e.g., performance outcomes) prominently upfront, while relegating visualizations of detailed hyperparameter manipulations or feature weight variations to later sections or the discussion, thus enhancing clarity for readers.

We thank the reviewer for pointing out the readability issue. We have corrected the x-axis label in Figure 7D. We hope this new layout with a simplified rule in Fig 3 and Fig 5 presents the key findings while retaining full mechanistic detail to make it easier to understand the model behavior.

iii. Writing clarity.

The authors have streamlined the "Metaplasticity" section and reduced references to dopamine, which is a positive step. However, the broader issue remains: the manuscript still appears overly detailed and more like a technical report of a novel learning rule, rather than a clearly structured scientific paper. I strongly recommend that the authors further distill the manuscript by clearly focusing on one or two central scientific questions or hypotheses-for instance, emphasizing core insights such as "inhibitory inputs facilitate nonlinear dendritic computations" or "distal dendritic inputs significantly contribute to nonlinear integration." Clarifying and highlighting these primary scientific questions early and consistently throughout the manuscript would substantially enhance readability and impact.

We appreciate the reviewer’s guidance on improving the manuscript’s clarity and focus.In response, we now highlight two central questions at the end of the Introduction and have retitled the main Results subsections to follow this thread, thereby sharpening the manuscript’s focus while retaining necessary technical detail (page3 line 20-28).We have also removed redundant passages and simplified technical details to improve overall readability .

Minor:

(1) The [Ca]NMDA in Figure 2A and 2C can have large values even when very few synapses are activated. Why is that? Is this setting biologically realistic?

The authors acknowledge that their simulated [Ca²⁺] levels exceed typical biological measurements but claim that the learning rule remains robust across variations in calcium concentrations. However, robustness to calcium variations was not explicitly demonstrated in the main figures. To convincingly address this concern, I recommend the authors explicitly test and present whether adopting biologically realistic calcium concentrations (~1 μM) impacts the learning outcomes or synaptic weight dynamics. Clarifying this point with a supplemental analysis or an additional figure panel would significantly strengthen their argument regarding the model's biological plausibility and robustness.

We thank the reviewer for the comment. The elevated [Ca²⁺]NMDA values reflect localized transients in spine heads with narrow necks and high NMDA conductance. These values are not problematic for our model, as the plasticity rule depends on relative calcium differences rather than absolute levels as the metaplasticity kernel will adjust. In future versions of our detailed neuron model, we will likely decrease the spine axial resistance of the spine neck.

Changes from the earlier version to the revised version

Abstract

Abstract

Added in the revised version: This study investigates the computational potential of single striatal projection neurons (SPN), emphasizing dendritic nonlinearities and their crucial role in solving complex integration problems. Utilizing a biophysically detailed multicompartmental model of an SPN, we introduce a calcium-based, local synaptic learning rule dependent on dendritic plateau potentials.

Earlier version: This study investigates the computational potential of single striatal projection neurons (SPN), emphasizing dendritic nonlinearities and their crucial role in solving complex integration problems. Utilizing a biophysically detailed multicompartmental model of an SPN, we introduce a calcium-based, local synaptic learning rule dependent on dendritic plateau potentials. According to what is known about excitatory corticostriatal synapses, the learning rule is governed by local calcium dynamics from NMDA and L-type calcium channels and dopaminergic reward signals. In order to devise a self-adjusting learning rule, which ensures stability for individual synaptic weights, metaplasticity is also used. We demonstrate that this rule allows single neurons to solve the nonlinear feature binding problem, a task traditionally attributed to neuronal networks. We also detail an inhibitory plasticity mechanism that contributes to dendritic compartmentalization, further enhancing computational efficiency in dendrites. This in silico study highlights the computational potential of single neurons, providing deeper insights into neuronal information processing and the mechanisms by which the brain executes complex computations.

Revised version: According to what is known about excitatory corticostriatal synapses, the learning rule is governed by local calcium dynamics from NMDA and L-type calcium channels and dopaminergic reward signals. In order to devise a self-adjusting learning rule, which ensures stability for individual synaptic weights, metaplasticity is also used. We demonstrate that this rule allows single neurons to solve the nonlinear feature binding problem, a task traditionally attributed to neuronal networks. We also detail an inhibitory plasticity mechanism that contributes to dendritic compartmentalization, further enhancing computational efficiency in dendrites. This in silico study highlights the computational potential of single neurons, providing deeper insights into neuronal information processing and the mechanisms by which the brain executes complex computations.

Introduction

Earlier version: Classically, single neurons in the nervous system have been thought to operate as simple linear integrators where the nonlinearity of dendrites can be neglected (McCulloch and Pitts, 1943). Based on this simplification, powerful artificial neural systems have been created, outperforming humans on multiple tasks (Silver et al., 2018). However, in recent decades it has been shown that active dendritic properties participate in shaping neuronal output, and also the dendrites can display nonlinear integration of input signals (Antic et al., 2010). Dendritic nonlinearities endow a neuron with the ability to perform sophisticated dendritic computations (Tran-Van-Minh et al., 2015; Gidon et al., 2020), expanding its computational power beyond what is available with the somatic voltage threshold and making it similar to a multilayer artificial neural network (Poirazi, Brannon and Mel, 2003).

Revised version: Classically, single neurons in the nervous system have been thought to operate as simple linear integrators such that the nonlinearity of dendrites can be neglected (McCulloch and Pitts, 1943). Based on this simplification, powerful artificial neural systems have been created that outperform humans on multiple tasks (Silver et al., 2018). However, in recent decades it has been shown that active dendritic properties shape neuronal output and that dendrites display nonlinear integration of input signals (Antic et al., 2010). These dendritic nonlinearities enable a neuron to perform sophisticated computations (Tran-Van-Minh et al., 2015; Gidon et al., 2020), expanding its computational power beyond what is available with the somatic voltage threshold and making it similar to a multilayer artificial neural network (Poirazi et al., 2003).

Earlier version: A dendritic nonlinearity common among projection neurons in several brain areas is the NMDA-dependent plateau potential (Oikonomou et al., 2014). Plateau potentials are regenerative, all-or-none, supralinear voltage elevations triggered by spatio-temporally clustered glutamatergic input (Schiller et al., 2000; Polsky, Mel and Schiller, 2004; Losonczy and Magee, 2006; Major et al., 2008; Larkum et al., 2009; Lavzin et al., 2012; Xu et al., 2012). Such plateaus require that nearby spines are coactivated, but the spatial requirement is somewhat loose as even single dendritic branches have been proposed to act as computational units (Losonczy and Magee, 2006; Branco and Häusser, 2010). With that said, multiple so-called hotspots, preferentially responsive to different input values or features, are known to form with close dendritic proximity (Jia et al., 2010; Chen et al., 2011; Varga et al., 2011). Such functional synaptic clusters are present in multiple species, developmental stages and brain regions (Kleindienst et al., 2011; Takahashi et al., 2012; Winnubst et al., 2015; Wilson et al., 2016; Iacaruso, Gasler and Hofer, 2017; Scholl, Wilson and Fitzpatrick, 2017; Niculescu et al., 2018; Kerlin et al., 2019; Ju et al., 2020). Hence, multiple features are commonly clustered in a single dendritic branch, indicating that this could be the neural substrate where combinations of simple features into more complex items occur.

Revised version: A dendritic nonlinearity common among projection neurons in several brain areas is the NMDA-dependent plateau potential (Oikonomou et al., 2014). Plateau potentials are regenerative, all-or-none, supralinear voltage elevations triggered by spatiotemporally clustered glutamatergic input (Schiller et al., 2000; Polsky et al., 2004; Losonczy and Magee, 2006; Major et al., 2008; Larkum et al., 2009; Lavzin et al., 2012; Xu et al., 2012). Such plateaus require that nearby spines are coactivated, but the spatial requirement is somewhat loose as even single dendritic branches have been proposed to act as computational units (Losonczy and Magee, 2006; Branco and Häusser, 2010). Nevertheless, multiple so-called hotspots, preferentially responsive to different input values or features, are known to form with close dendritic proximity (Jia et al., 2010; Chen et al., 2011; Varga et al., 2011). Such functional synaptic clusters are present in multiple species, developmental stages and brain regions (Kleindienst et al., 2011; Takahashi et al., 2012; Winnubst et al., 2015; Wilson et al., 2016; Iacaruso et al., 2017; Scholl et al., 2017; Niculescu et al., 2018; Kerlin et al., 2019; Ju et al., 2020). Hence, multiple features are commonly clustered in a single dendritic branch, indicating that this could be the neural substrate where combinations of simple features into more complex items occur.

Earlier version: Combinations of features in dendritic branches further provide single neurons with the possibility to solve linearly non-separable tasks, such as the nonlinear feature binding problem (NFBP) (Tran-Van-Minh et al., 2015; Gidon et al., 2020). In its most basic form, the NFBP involves discriminating between two groups of feature combinations. This problem is nonlinear because the neuron must learn to respond only to specific feature combinations, even though all features contribute equally in terms of synaptic input. A commonly used example involves two different shapes combined with two different colors, resulting in four total combinations. Out of these, the neuron should respond to only two specific feature combinations (exemplified in Figure 1A and 1B).

Revised version: Combinations of features in dendritic branches further provide single neurons with the possibility to solve linearly non-separable tasks, such as the nonlinear feature binding problem (NFBP) (Tran-Van-Minh et al., 2015; Gidon et al., 2020). In its most basic form, the NFBP involves discriminating between two groups of feature combinations. This problem is nonlinear because the neuron must learn to respond only to specific feature combinations, even though all features contribute equally in terms of synaptic input. A commonly used example involves two different shapes combined with two different colors, resulting in four total combinations. Out of these, the neuron should respond only to two specific feature combinations (exemplified in Figure 1A and B).

Learning Mechanisms in direct pathway Striatal Projection Neurons (dSPNs) for the Nonlinear Feature Binding Problem (NFBP)

Earlier version: (A) : Inputs and assumed supralinearity that could solve the NFBP: The NFBP is represented with an example from visual feature binding. In the simplest form of the NFBP, a stimulus has two features, here shape and form, each with two possible values, strawberry and banana, and red and yellow, respectively. In the NFBP the neuron should learn to respond by spiking to two of the feature combinations, representing the relevant stimuli (red strawberry and yellow banana), while remaining silent for the other two feature combinations which represent the irrelevant stimuli (yellow strawberry and red banana). Assuming that each feature is represented with locally clustered synapses, a solution of the NFBP can be achieved when the co-active clusters on a single dendrite, corresponding to a relevant stimulus, evoke a plateau potential, thus superlinearly exciting the soma. Conversely, co-activation of synaptic clusters for the irrelevant combinations should not evoke plateau potentials.

Revised version: (A) : Inputs and assumed supralinearity that could solve the NFBP: The NFBP is represented with an example from visual feature binding. In the simplest form of the NFBP, a stimulus has two features, here shape and form, each with two possible values, strawberry and banana, and red and yellow, respectively. In the NFBP the neuron should learn to respond by spiking to two of the feature combinations, representing the relevant stimuli (red strawberry and yellow banana), while remaining silent for the other two feature combinations which represent the irrelevant stimuli (yellow strawberry and red banana). Assuming that each feature is represented with locally clustered synapses, a solution of the NFBP can be achieved when the co-active clusters on a single dendrite, corresponding to a relevant stimulus, evoke a plateau potential, thus superlinearly exciting the soma. Conversely, co-activation of synaptic clusters for the irrelevant combinations should not evoke plateau potentials. (B) : Synaptic clustering in dendrites: Illustration of how synaptic plasticity in SPNs may contribute to solving the NFBP for a pre-existing arrangement of synaptic clusters on two dendrites. A plasticity rule which strengthens only synaptic clusters representing relevant feature combinations, so that they produce robust supralinear responses, while weakening synapses activated by irrelevant feature combinations, could solve the NFBP. (C) : Dopamine (Da) feedback: dopaminergic inputs from the midbrain to the striatum (Str) guides the learning process, differentiating between positive feedback for relevant stimuli and negative feedback for irrelevant stimuli. Positive feedback represented by dopamine peaks is necessary for LTP, and negative feedback represented by a dopamine pause is necessary for LTD. (D) : Signaling pathways underlying synaptic plasticity in dSPNs: Illustrations of signaling components at the corticostriatal synapse that modify synaptic strength (redrawn from Shen et al., 2008). NMDA calcium influx, followed by stimulation of D1 dopamine receptors (D1Rs), triggers LTP (while inhibiting the LTD cascade). L-type calcium influx and activation of metabotropic glutamate receptors (mGluRs) when D1Rs are free of dopamine triggers LTD (while counteracting the LTP cascade).

Removed in the revised version: (B) : Synaptic clustering in dendrites: Illustration of how synaptic plasticity in SPNs may contribute to solving the NFBP for a pre-existing arrangement of synaptic clusters on two dendrites. A plasticity rule which strengthens only synaptic clusters representing relevant feature combinations, so that they produce robust supralinear responses, while weakening synapses activated by irrelevant feature combinations, could solve the NFBP.

Removed in the revised version: (C) : Dopamine (Da) feedback: dopaminergic inputs from the midbrain to the striatum (Str) guides the learning process, differentiating between positive feedback for relevant stimuli and negative feedback for irrelevant stimuli. Positive feedback represented by dopamine peaks is necessary for LTP, and negative feedback represented by a dopamine pause is necessary for LTD.

Removed in the revised version: (D) : Signaling pathways underlying synaptic plasticity in dSPNs: Illustrations of signaling components at the corticostriatal synapse that modify synaptic strength (redrawn from Shen et al., 2008). NMDA calcium influx, followed by stimulation of D1 dopamine receptors (D1Rs), triggers LTP (while inhibiting the LTD cascade). L-type calcium influx and activation of metabotropic glutamate receptors (mGluRs) when D1Rs are free of dopamine triggers LTD (while counteracting the LTP cascade).

Earlier version: As a task, the NFBP is relevant to brain regions which perform integration of multimodal input signals, or signals representing different features of the same modality (Roskies, 1999). It is usually illustrated with examples from the visual system, as in Figure 1A (Roskies, 1999; von der Malsburg, 1999; Tran-Van-Minh et al., 2015). A region that integrates multimodal inputs, such as sensory information and motor-related signals, is the input nucleus of the basal ganglia, the striatum (Reig and Silberberg, 2014; Johansson and Silberberg, 2020), and this system will be used in the present modeling study. We will here, however, continue to illustrate the NFBP with the more intuitive features borrowed from the visual field, although for the dorsal striatum these features would rather map onto different sensory- and motor-related features. Plateau potentials and some clustering of input have been demonstrated in SPNs (Plotkin, Day and Surmeier, 2011; Oikonomou et al., 2014; Du et al., 2017; Hwang et al., 2022; Day et al., 2024; Sanabria et al., 2024).

Revised version: As a task, the NFBP is relevant to brain regions which perform integration of multimodal input signals, or signals representing different features of the same modality (Roskies, 1999). It is usually illustrated with examples from the visual system, as in Figure 1A (Roskies, 1999; von der Malsburg, 1999; Tran-Van-Minh et al., 2015). A region that integrates multimodal inputs, such as sensory information and motor-related signals, is the input nucleus of the basal ganglia, the striatum (Reig and Silberberg, 2014; Johansson and Silberberg, 2020), and this system will be used in the present modeling study. Here, however, we will continue to illustrate the NFBP with the more intuitive features borrowed from the visual field, although for the dorsal striatum these features would rather map onto different sensory- and motor-related features. Plateau potentials and some clustering of input have been demonstrated in SPNs (Plotkin et al., 2011; Oikonomou et al., 2014; Du et al., 2017; Hwang et al., 2022; Day et al., 2024; Sanabria et al., 2024).

Earlier version: In addition to integrating converging input from the cortex and the thalamus, the striatum is densely innervated by midbrain dopaminergic neurons which carry information about rewarding stimuli (Schultz, 2007; Matsuda et al., 2009; Surmeier et al., 2010). As such, the striatum is thought to be an important site of reward learning, associating actions with outcomes based on neuromodulatory cues. In this classical framework, peaks in dopamine signify rewarding outcomes and pauses in dopamine represent the omission of expected rewards (Schultz, Dayan and Montague, 1997). Dopamine signals further control the synaptic plasticity of corticostriatal synapses on the SPNs (Figure 1C). In direct pathway SPNs (dSPN) expressing the D1 receptor, a dopamine peak together with significant calcium influx through NMDA receptors triggers synaptic strengthening (long-term potentiation - LTP). Conversely, when little or no dopamine is bound to the D1 receptors, as during a dopamine pause, and there is significant calcium influx through L-type calcium channels, synaptic weakening occurs (long-term depression - LTD, see Figure 1D) (Shen et al., 2008; Fino et al., 2010; Plotkin et al., 2013).

Revised version: In addition to integrating converging input from the cortex and the thalamus, the striatum is densely innervated by midbrain dopaminergic neurons which carry information about rewarding stimuli (Schultz, 2007; Matsuda et al., 2009; Surmeier et al., 2010). As such, the striatum is thought to be an important site of reward learning, associating actions with outcomes based on neuromodulatory cues. In this classical framework, peaks in dopamine signify rewarding outcomes and pauses in dopamine represent the omission of expected rewards (Schultz et al., 1997). Dopamine signals further control the synaptic plasticity of corticostriatal synapses on the SPNs (Figure 1C). In direct pathway SPNs (dSPN) expressing the D1 receptor, a dopamine peak together with significant calcium influx through NMDA receptors triggers synaptic strengthening (long-term potentiation – LTP). Conversely, when little or no dopamine is bound to the D1 receptors, as during a dopamine pause, and there is significant calcium influx through L-type calcium channels, synaptic weakening occurs (long-term depression – LTD, see Figure 1D) (Shen et al., 2008; Fino et al., 2010; Plotkin et al., 2013).

Earlier version: The ability to undergo LTP or LTD is itself regulated (Huang et al., 1992), a concept termed metaplasticity (Abraham and Bear, 1996). Metaplasticity refers to changes in synaptic plasticity driven by prior synaptic activity (Frey et al., 1995) or by neuromodulators (Moody, Carlisle and O’Dell, 1999), effectively making plasticity itself adaptable. Metaplasticity can further regulate synaptic physiology to shape future plasticity without directly altering synaptic efficacy, acting as a homeostatic mechanism to keep synapses within an optimal dynamic range (Abraham, 2008).

Revised version: The ability to undergo LTP or LTD is itself regulated (Huang et al., 1992), a concept termed metaplasticity (Abraham and Bear, 1996). Metaplasticity refers to changes in synaptic plasticity driven by prior synaptic activity (Frey et al., 1995) or by neuromodulators (Moody et al., 1999), effectively making plasticity itself adaptable. Metaplasticity can further regulate synaptic physiology to shape future plasticity without directly altering synaptic efficacy, acting as a homeostatic mechanism to keep synapses within an optimal dynamic range (Abraham, 2008).Previous theoretical studies (Bienenstock et al., 1982; Fusi et al., 2005; Clopath et al., 2010; Benna and Fusi, 2016; Zenke and Gerstner, 2017) demonstrate the essential role of metaplasticity in maintaining stability in synaptic weight distributions.

Earlier version: If dopamine peaks are associated with the relevant feature combinations in the NFBP and dopamine pauses with the irrelevant ones, it can trigger LTP in synapses representing the relevant feature combinations and LTD in those representing irrelevant combinations. If, after learning, the relevant feature combinations have strong enough synapses so they can evoke plateau potentials while the irrelevant feature combinations have weak enough synapses so they don’t evoke plateaus, the outcome of this learning process should be a synaptic arrangement that could solve the NFBP (Figure 1B) (Tran-Van-Minh et al., 2015). In line with this, it has been demonstrated that the NFBP can be solved in abstract neuron models where the soma and dendrites are represented by single electrical compartments and where neuronal firing and plateau potentials are phenomenologically represented by instantaneous firing rate functions (Legenstein and Maass, 2011; Schiess, Urbanczik and Senn, 2016). Good performance on the NFBP has also been demonstrated with biologically detailed models (Bicknell and Häusser, 2021). This solution used a multicompartmental model of a single pyramidal neuron, including both excitatory and inhibitory synapses and supralinear NMDA depolarizations. Synapses representing different features were randomly dispersed throughout the dendrites and a phenomenological precalculated learning rule - dependent on somatic spike timing and high local dendritic voltage - was used to optimize the strength of the synapses. The solution did, however, depend on a form of supervised learning as somatic current injections were used to raise the spiking probability of the relevant feature combinations.

Revised version: If dopamine peaks are associated with the relevant feature combinations in the NFBP and dopamine pauses with the irrelevant ones, they trigger LTP in synapses representing the relevant feature combinations and LTD in those representing irrelevant combinations. If, after learning, the relevant feature combinations have strong enough synapses so they can evoke plateau potentials while the irrelevant feature combinations have weak enough synapses so they don’t evoke plateaus, the outcome of this learning process should be a synaptic arrangement that could solve the NFBP (Figure 1B) (Tran-Van-Minh et al., 2015). In line with this, it has been demonstrated that the NFBP can be solved in abstract neuron models where the soma and dendrites are represented by single electrical compartments and where neuronal firing and plateau potentials are phenomenologically represented by instantaneous firing rate functions (Legenstein and Maass, 2011; Schiess et al., 2016). Good performance on the NFBP has also been demonstrated with biologically detailed models (Bicknell and Häusser, 2021). This solution used a multicompartmental model of a single pyramidal neuron, including both excitatory and inhibitory synapses and supralinear NMDA depolarizations. Synapses representing different features were randomly dispersed throughout the dendrites and a phenomenological precalculated learning rule – dependent on somatic spike timing and high local dendritic voltage – was used to optimize the strength of the synapses. The solution did, however, depend on a form of supervised learning as somatic current injections were used to raise the spiking probability of the relevant feature combinations.

Added in the revised version: In this study, we ask whether—and under what conditions—the theoretical solution to the nonlinear feature-binding problem (NFBP) can be achieved in a biophysically detailed model of a striatal projection neuron (SPN) using only local, biologically grounded learning rules. We frame this paper around two questions: First, can a single dSPN equipped with only calcium- and dopaminedependent excitatory plasticity solve the NFBP when the relevant features are pre-clustered on one dendritic branch? Second, if that mechanism is insufficient—as with randomly distributed or very distal inputs—does adding branch-specific inhibitory plasticity restore plateau-based nonlinear computation and spiking?

Added in the revised version: To answer these questions we adopt an approach that relies on the following key mechanisms:

Added in the revised version: A local learning rule

Removed in the revised version: In this study, we aim to investigate whether and to what extent the theoretical solution to the NFBP, as proposed by e.g. Tran-Van-Minh et al., (2015) can be achieved using biophysically detailed models of SPNs and a biologically plausible local learning rule. Our approach relies on the following key mechanisms:

Earlier version: A local learning rule: We develop a learning rule driven by local calcium dynamics in the synapse and by reward signals from the neuromodulator dopamine. This plasticity rule is based on the known synaptic machinery for triggering LTP or LTD in the corticostriatal synapse onto dSPNs (Shen et al., 2008). Importantly, the rule does not rely on supervised learning paradigms and neither is a separate training and testing phase needed.

Revised version: We develop a learning rule driven by local calcium dynamics in the synapse and by reward signals from the neuromodulator dopamine. This plasticity rule is based on the known synaptic machinery for triggering LTP or LTD at the corticostriatal synapse onto dSPNs (Shen et al., 2008). Importantly, the rule does not rely on supervised learning paradigms, and no separate training and testing phase is required.

Added in the revised version: Robust dendritic nonlinearities

Earlier version: Robust dendritic nonlinearities: According to Tran-Van-Minh et al., (2015) sufficient supralinear integration is needed to ensure that e.g. two inputs (i.e. one feature combination in the NFBP, Figure 1A) on the same dendrite generate greater somatic depolarization than if those inputs were distributed across different dendrites. To accomplish this we generate sufficiently robust dendritic plateau potentials using the approach in Trpevski et al., (2023).

Revised version: According to Tran-Van-Minh et al. (2015), sufficient supralinear integration is needed to ensure that, e.g., two inputs (one feature combination in the NFBP, Figure 1A) on the same dendrite generate greater somatic depolarization than if those inputs were distributed across different dendrites. To accomplish this, we generate sufficiently robust dendritic plateau potentials using the approach of Trpevski et al. (2023).

Added in the revised version: Metaplasticity

Added in the revised version: Our simulations demonstrate that metaplasticity is necessary for synaptic weights to remain stable and within physiologically realistic ranges, regardless of their initial values.

Removed in the revised version: Metaplasticity: Although not discussed much in more theoretical work, our study demonstrates the necessity of metaplasticity for achieving stable and physiologically realistic synaptic weights. This mechanism ensures that synaptic strengths remain within biologically plausible ranges during training, regardless of initial synaptic weights.

We first demonstrate the effectiveness of the proposed learning rule under the assumption of pre-existing clustered synapses for each individual feature, as suggested by Tran-Van-Minh et al. (2015). These clustered synapses are trained, to a degree where they can reliably evoke robust plateau potentials for the relevant feature combinations required to solve the NFBP, while synapses representing irrelevant features are weakened (illustrated in Figure 1B).

We then extend the analysis by applying the learning rule to more randomly distributed synapses, which initially exhibit minimal local supralinear integration. However, when incorporating the assumption that branch-specific plasticity mechanisms are at play, supralinear integration emerges within distinct dendritic branches. This suggests that branch-specific plasticity could play a critical role in enabling single neurons to solve nonlinear problems. Furthermore, we explore an activity-dependent rule for GABAergic plasticity and demonstrate its potential importance in shaping dendritic nonlinearities. This mechanism may thus further enhance computational capabilities by refining and stabilizing the integration of inputs across dendritic branches.

Although brain systems like the striatum, which integrate multimodal inputs, somehow solve nonlinear problems at the network or systems level, it remains unclear whether individual neurons in the brain regularly solve the NFBP. Our investigation suggests, however, that single SPNs possess the computational capacity to address linearly non-separable tasks. This is achieved by leveraging the organism’s performance feedback, represented by dopamine peaks (success) and dopamine pauses (failure), in combination with their ability to generate robust dendritic plateau potentials. Since the mechanisms used in the rule are general to the brain, this capability may also extend to other projection neurons capable of producing dendritic plateaus, such as pyramidal neurons. However, the specific feedback mechanisms, represented here by dopamine, would need to be associated with alternative neuromodulatory signals depending on the type of neuron and synapse involved.

Results

Characterization of the dendritic nonlinearities in the model

Earlier version: The nonlinear, sigmoidal voltage sensitivity of the NMDA receptors is a crucial element for forming dendritic plateau potentials. We use a model for generating plateau potentials first presented in Gao et al., (2021) and adjusted to SPNs (Trpevski et al., 2023). To produce robust, all-or-none plateau potentials, glutamate spillover from the synaptic cleft that activates extrasynaptic NMDA receptors is included in the model. Glutamate spillover occurs when the total synaptic weight of the nearby activated synapses reaches a threshold value (see Methods). The threshold value here is set to be equivalent to the total weight of 8 clustered synapses with weights of 0.25 each (weights of 0.25 corresponds to a maximal conductance of 0.625 nS). Figure 2A shows the somatic membrane potential following synaptic activation of a cluster of synapses of increasing size and the corresponding local spine membrane potential as well as the NMDA and L-type calcium accumulated in a single spine (averaged over all spines in the cluster). A plateau potential is generated when a critical level of total NMDA conductance in a dendritic segment is reached (accomplished here by the addition of more synapses in a cluster eventually leading to glutamate spillover). Reaching the spillover threshold produces a sudden and prolonged increase in NMDA conductance, caused by the activation of extrasynaptic NMDA receptors (where the clearance of glutamate is assumed to be slower). The result is a more robust plateau potential that better resembles experimentally generated plateau potentials, as investigated in more detail in Trpevski et al., (2023).

Revised version: The nonlinear sigmoidal voltage sensitivity of NMDA receptors is a crucial element for forming dendritic plateau potentials. We used a model to generate plateau potentials first presented in Gao et al. (2021) and adjusted to SPNs (Trpevski et al., 2023). To produce robust all-or-none plateau potentials, the model includes glutamate spillover from the synaptic cleft that activates extrasynaptic NMDA receptors. Glutamate spillover occurs when the total synaptic weight of the nearby activated synapses reaches a threshold value (see Methods). The threshold value here is set to be equivalent to the total weight of eight clustered synapses with weights of 0.25 each (weights of 0.25 correspond to a maximal conductance of 0.625 nS).

Added in the revised version: Figure 2A shows the somatic membrane potential following synaptic activation of a cluster of synapses of increasing size and the corresponding local spine membrane potential as well as the NMDA and L-type calcium accumulated in a single spine (averaged over all spines in the cluster). A plateau potential is generated when a critical level of total NMDA conductance in a dendritic segment is reached (accomplished here by the addition of more synapses in a cluster eventually leading to glutamate spillover). Reaching the spillover threshold produces a sudden and prolonged increase in NMDA conductance, caused by the activation of extrasynaptic NMDA receptors (where the clearance of glutamate is assumed to be slower). The result is a more robust plateau potential that better resembles experimentally generated plateau potentials, as investigated in more detail in Trpevski et al. (2023).

Characterization of dendritic plateau potentials in the model.

Removed in the revised version: (A) : Somatic voltage, spine voltage, NMDA calcium, and L-type calcium evoked by a cluster varying in size from 1 to 20 synapses. A plateau potential is evoked when glutamate spillover is activated, here triggered when 8 synapses with a weight of 0.25 each are coactivated (corresponding to the “baseline” weights in C).

Removed in the revised version: (B) : Schematic of the neuron morphology with an arrow indicating the stimulated dendritic branch shown in panel A.

Earlier version: (C) : The mean maximal amplitude, with standard deviation shown in bars, of the measures in A averaged over 10 different dendrites, and 10 trials per dendrite. The curves represent clusters with different synaptic weights: baseline (0.25, 0.625 nS), strengthened (0.4, 1 nS) and weakened weights (0.2, 0.5 nS). Synaptic background noise is used in all simulations to elevate the membrane potential to ranges seen in vivo (Reig and Silberberg, 2014).

Revised version: (A) : Somatic voltage, spine voltage, NMDA calcium, and L-type calcium evoked by a cluster varying in size from 1 to 20 synapses. A plateau potential is evoked when glutamate spillover is activated, here triggered when 8 synapses with a weight of 0.25 each are coactivated (corresponding to the “baseline” weights in C). (B) : Schematic of the neuron morphology with an arrow indicating the stimulated dendritic branch shown in panel A. (C) : The mean maximal amplitude, with standard deviation shown in bars, of the measures in A averaged over 10 different dendrites, and 10 trials per dendrite. The curves represent clusters with different synaptic weights: baseline (0.25, 0.625 nS), strengthened (0.4, 1 nS) and weakened weights (0.2, 0.5 nS). Synaptic background noise is used in all simulations to elevate the membrane potential to ranges seen in vivo (Reig and Silberberg, 2014).

Earlier version: The synaptic input to the neuron is provided through the activation of a cluster of synapses at the indicated location in Figure 2B, and gives the voltage and calcium responses as exemplified in Figure 2A. Figure 2C shows the maximal amplitudes of the somatic and spine voltage together with NMDA and L-type calcium signals in the dendritic spines, averaged over 10 trials and over 10 dendrites. The “baseline” results in Figure 2C are within the range of the initial synaptic weights of excitatory synapses in all remaining figures in the article, and thus illustrate a possible initial situation before learning. Stronger and weaker synapses require a smaller and a larger cluster to trigger a plateau potential, respectively (Figure 2C, spine voltage panel). In the simulation using “strengthened” synapses the synaptic weights in the cluster are 60% greater than in the “baseline” case, and hence need fewer synapses to trigger plateau potentials. Conversely, with weaker synapses, where weights are 25% smaller than the “baseline” case, more synapses are needed to evoke a plateau.

Revised version: The synaptic input to the neuron is provided through the activation of a cluster of synapses at the indicated location in Figure 2B, and gives the voltage and calcium responses as exemplified in Figure 2A. Figure 2C shows the maximal amplitudes of the somatic and spine voltage together with NMDA and L-type calcium signals in the dendritic spines, averaged over 10 trials and over 10 dendrites. The “baseline” results in Figure 2C are within the range of the initial synaptic weights of excitatory synapses in all remaining figures in the article, and thus illustrate a possible initial situation before learning. Stronger and weaker synapses require a smaller and a larger cluster to trigger a plateau potential, respectively (Figure 2C, spine voltage panel). In the simulation using “strengthened” synapses the synaptic weights in the cluster are 60% greater than in the “baseline” case, and hence need fewer synapses to trigger plateau potentials. Conversely, with weaker synapses, where weights are 20% smaller than the “baseline” case, more synapses are needed to evoke a plateau.

To summarize, the dSPN model exhibits the dendritic nonlinearities required for solving the NFBP. Further, clusters of strengthened synapses can reliably generate robust plateau potentials, which produces a long lasting somatic depolarization and thereby increases the likelihood for somatic spiking (in accordance with (Du et al., 2017) and also illustrated in (Trpevski et al., 2023)). Conversely, clusters of weakened synapses will most likely not generate plateau potentials, and thus the neuron will spike with much lower probability following activation of such a cluster.

Characterization of the synaptic plasticity rule

To characterize the learning rule we started with a simple setup where three features, each illustrating either a color or a shape, were distributed onto two dendritic branches. The setup was such that one relevant and one irrelevant feature combination were represented in each dendrite and the relevant–irrelevant feature combination was unique to each dendrite (see illustration in Figure 3A). Each feature was represented with 5 synapses, and we start with the assumption that those synapses are already organized in pre-existing clusters. In addition to background synaptic noise inputs, 108 randomly distributed glutamatergic synapses were also added. These synapses were activated together with all four stimuli, i.e. they were feature-unspecific. SPNs have a very hyperpolarized resting potential, and the additional feature-unspecific synapses allowed the neurons to spike often enough in the beginning of the learning process so that a dopamine feedback signal would be elicited and trigger learning in the activated synapses.

Example of setup and learning-induced synaptic plasticity

Removed in the revised version: (A) : Illustration of input configuration. The upper panel illustrates the arrangement of features across two dendrites. Each dendrite has synaptic clusters for three features, allowing the representation of each of the four stimuli if seen from the whole neuron’s perspective (but only one relevant stimulus per dendrite). The bottom panel illustrates the stimulation protocol used in the simulation, where stimulus presentation is followed by a dopaminergic feedback signal only if the neuron spikes. Two stimuli are spaced 800 ms apart, to allow for the calcium dynamics to reach baseline levels.

Removed in the revised version: (B) : Example voltage in the soma and the middle of dendrite 1 (d1) and dendrite 2 (d2) before and after learning. Each dendrite here stops responding to the respective irrelevant stimuli during learning.

Earlier version: (C), (D): Evolution of synaptic conductances during learning (top row) and examples of peak Ca²⁺ levels (dots) alongside kernel dynamics in single synapses (three lower rows). Panels in (C) depict clustered synapses in dendrite 1 (d1), where ‘yellow’ and ‘banana’ generally undergo LTP, while ‘red’ undergoes LTD. Panels in (D) show distributed, feature-unspecific synapses. Among these, the purple and blue traces represent synapses that are weakened, while the green trace exemplifies a synapse near one of the clusters that, by chance, achieves a sufficiently high local NMDA calcium level for LTP to dominate. The initial synaptic conductances are drawn from a normal distribution with a mean of 0.625 and a standard deviation of 0.375 nS. MP-K and LTP-K refer to metaplasticity and LTP kernels, respectively, while the solid line represents the midpoint of the kernels, where LTP is strongest. ‘Max’ indicates the peak NMDA calcium during a single stimulus.

Revised version: (A): Input configuration. The panel shows how the four stimulus features are distributed across two dendrites (d1, d2). Each dendrite contains pre-existing synaptic clusters for three features (black circles), and distributed, feature-unspecific synapses (shown in grey). (B) : Stimulation timeline and plasticity schematic. The top diagram depicts a 20 ms stimulus followed by a 50 ms dopaminergic feedback pulse delivered with a 300 ms delay, only if the neuron spikes. Successive stimuli are separated by 800 ms to let calcium return to baseline. Lower panels: a dopamine peak (left) gates LTP—synapses whose [Ca]NMDA falls within the bell-shaped window are potentiated (w ↑, green) while the kernel midpoint θLTP shifts downward (↓, red); a dopamine pause (right) gates LTD—the degree of depression scales with [Ca]L-type (w ↓, red) and θLTP shifts upward (↑, green). (C), (D): Evolution of synaptic conductances during learning (top row) and examples of peak Ca2+ levels (black dots) alongside kernel dynamics in single synapses (three lower rows). Panels in (C) depict clustered synapses in dendrite 1 (d1), where ‘yellow’ and ‘banana’ generally undergo LTP, while ‘red’ undergoes LTD. See Figure 3—video 1for an animation of panel C showing the first 400 training stimuli of the learning sequence. Panels in (D) show distributed, feature-unspecific synapses. Among these, the Example 1 and Example 3 traces represent synapses that are weakened, while the Example 2 trace exemplifies a synapse near one of the clusters that, by chance, achieves a sufficiently high local NMDA calcium level for LTP to dominate. The initial synaptic conductances are drawn from a normal distribution with a mean of 0.625 and a standard deviation of 0.125 nS. The solid line represents the midpoint of the kernels, where LTP is strongest. ‘Max’ indicates the peak NMDA calcium during a single stimulus. The darker shaded regions represent the LTP kernel and the lighter shaded ones show the wider metaplasticity kernel. (E): Example voltage in the soma and the middle of dendrite 1 (d1) and dendrite 2 (d2) before and after learning. Each dendrite here stops responding to the respective irrelevant stimuli during learning.

Added in the revised version: Learning-induced synaptic plasticity with metaplasticity turned off

Added in the revised version: (A): Example voltage traces recorded at the soma and in two dendrites before and after learning. Without metaplasticity, no significant learning occurs—every stimulus triggers spikes (50% performance). (B, C): Evolution of synaptic conductances throughout learning (top row) and characteristic examples of peak calcium levels (dots), along with the constant kernels in single synapses (shaded areas, bottom rows). The same synaptic arrangement and initial synaptic weights as in Figure 3A is used (for both clustered and feature-unspecific synapses). (B) illustrates changes in clustered synapses on dendrite 1 (d1). (C) shows distributed, feature-unspecific synapses. The final outcome depends on the initial synaptic weights. Synapses whose calcium levels are within or above the LTP region cannot be weakened, because any LTD during irrelevant stimuli is counteracted by LTP occurring for the relevant stimuli. The calcium levels remain “trapped” within the LTP kernel, and the synaptic weights “zig-zag” around a steady state level (seen in the example clustered ‘red’, ‘yellow’ and ‘banana’ synapse, as well as for the Example 2 and Example 3 feature-unspecific synapses). The synapses with small initial weights, whose calcium levels are situated below the LTP kernel are weakened (Example 1 synapse). The solid lines within the LTP kernels represent the midpoint of the kernels where LTP is strongest. In summary, without metaplasticity learning of NFBP is difficult if initial synaptic weights are not optimally tuned with regard to the postsynaptic calcium concentration. Metaplasticity allows the synapses to fall in and out of the LTP kernel range. They can be strengthened as long as their calcium levels are within the LTP kernel, are stabilized when they are above the LTP kernel, and can be weakened when they are below the LTP kernel.

Earlier version: The learning rule modified synaptic weights based on local calcium levels at each synapse, employing distinct plasticity kernels for LTP and LTD, each associated with different calcium sources—NMDA and L-type channels, respectively (Figure 1; see also Methods). The LTD kernel was governed by a calcium threshold, above which the change of synaptic weight was proportional to the amplitude of L-type calcium. That is, higher calcium levels trigger a larger reduction of weights than low concentrations. Conversely, the LTP kernel was represented by a bell-shaped function over a range of NMDA calcium concentrations, allowing LTP in the presence of dopamine when calcium levels fell within this range, while calcium outside this range did not elicit LTP (see Methods).

Revised version: The learning rule modified synaptic weights based on local calcium levels at each synapse, employing distinct plasticity kernels for LTP and LTD, each associated with different calcium sources—NMDA and L-type channels, respectively (Figure 1; schematic in Figure 3B; see also Methods). The LTD kernel was governed by a calcium threshold, above which the change of synaptic weight was proportional to the amplitude of L-type calcium. That is, higher calcium levels trigger a larger reduction of weights than low concentrations. Conversely, the LTP kernel was represented by a bell-shaped function over a range of NMDA calcium concentrations, allowing LTP in the presence of dopamine when calcium levels fell within this range, while calcium outside this range did not elicit LTP (see Methods).

Earlier version: An additional feature of the learning rule was metaplasticity (Abraham, 2008), which dynamically adjusted the calcium dependence of the LTP kernel. That is, the LTP kernel’s calcium dependence was adjusted over time (see Methods). Metaplasticity enabled e.g. previously weak synapses that were repeatedly co-activated with rewards, to eventually participate in LTP, as their calcium threshold was gradually lowered when the LTP-kernel was shifted. The same mechanism also prevented excessive strengthening of already potent synapses as also the upper calcium threshold for LTP was lowered following repeated co-activation with rewards. This behavior of the LTP kernel ensured stability in synaptic weights and allowed for continuous refinement of synaptic efficacy based on activation pattern and presence/absence of dopamine.

Revised version: An additional feature of the learning rule was metaplasticity (Abraham, 2008), which dynamically adjusted the calcium dependence of the LTP kernel. That is, the LTP kernel’s calcium dependence was adjusted over time (see Figure 3B–D). Metaplasticity enabled, e.g., previously weak synapses that were repeatedly co-activated with rewards to eventually participate in LTP, as their calcium threshold was gradually lowered when the LTP kernel was shifted. The same mechanism also prevented excessive strengthening of already potent synapses as shifting the LTP kernel downwards also reduced the potentiation of synapses with large calcium transients. This behavior of the LTP kernel ensured stability in synaptic weights and allowed for continuous refinement of synaptic efficacy based on activation pattern and presence/absence of dopamine.

Earlier version: Initially, the synaptic weight of each synapse was set to a value chosen uniformly at random from a range of 0.25 ± 0.15 (corresponding to conductances of 0.625 ± 0.375 nS). Initially, the synapses therefore experienced different NMDA and L-type calcium levels following activation, and hence their weights were also updated differently based on where the calcium level fell within their individual plasticity kernels.

Revised version: Initially, the synaptic weight of each synapse was set to a value chosen uniformly at random from a range of 0.25±0.05 (corresponding to conductances of 0.625±0.125 nS). Initially, the synapses therefore experienced different NMDA and L-type calcium levels following activation, and hence their weights were also updated differently based on where the calcium level fell within their individual plasticity kernels.

Earlier version: The neuron model was during training activated with a sequence of 960 feature combinations, including equal amounts of relevant (i.e., ‘red strawberry’ and ‘yellow banana’) and irrelevant feature combinations (i.e., ‘yellow strawberry’ and ‘red banana’). Thus, dopamine peaks and dopamine pauses occurred equally often at the beginning of the learning phase. When the neuron spiked for the relevant feature combinations, dopamine rewards were delivered, triggering LTP in the active synapses with NMDA Ca levels within the LTP kernel. Conversely, spiking for the irrelevant feature combinations elicited a dopamine pause as feedback, triggering LTD as a function of L-type Ca. As learning progressed, the neuron learned to respond only to the relevant feature combinations (see example in Figure 3B). Initially, all four stimuli—‘yellow banana’ and ‘red banana’ in dendrite 1, and ‘red strawberry’ and ‘yellow strawberry’ in dendrite 2—elicited robust supralinear responses, as they together reached the threshold for glutamate spillover in the model. This is shown with gray lines in the voltage traces in Figure 3B. After learning, the neuron could differentiate between the two sets of stimuli. The relevant feature combinations associated with a reward continued to provoke a plateau potential and eliciting somatic spiking, while in contrast, the neuron’s response to irrelevant feature combinations was notably decreased (shown with black lines in the somatic voltage trace in Figure 3B).

Revised version: During training, the neuron model was activated with a sequence of 1400 feature combinations, including equal amounts of relevant (i.e. ‘red strawberry’ and ‘yellow banana’) and irrelevant feature combinations (i.e. ‘yellow strawberry’ and ‘red banana’). Thus, dopamine peaks and dopamine pauses occurred equally often at the beginning of the learning phase. When the neuron spiked for the relevant feature combinations, dopamine rewards were delivered, triggering LTP in the active synapses with NMDA calcium levels within the LTP kernel. Conversely, spiking for the irrelevant feature combinations elicited a dopamine pause as feedback, triggering LTD as a function of L-type calcium. Initially, all four stimuli—’yellow banana’ and ‘red banana’ in dendrite 1, and ‘red strawberry’ and ‘yellow strawberry’ in dendrite 2—elicited robust supralinear responses, as they together reached the threshold for glutamate spillover in the model (Figure 3E, gray lines). After learning, however, the neuron could differentiate between the two sets of stimuli—the relevant feature combinations associated with a reward continued to evoke a plateau potential and elicit somatic spiking, while in contrast, the neuron’s response to irrelevant feature combinations was notably decreased (Figure 3E, black lines).

Earlier version: Figure 3C shows the evolution of synaptic conductances during the learning process for dendrite 1 (dendrite 2 is not shown, but has a similar behavior for relevant and irrelevant stimuli). The synapses representing the relevant feature combination in this dendrite (‘yellow’ and ‘banana’) are typically strengthened, eventually encoding this stimulus robustly. Conversely, the synapses for the feature ‘red’, activated during the irrelevant feature combination (‘red banana’) are all weakened, making the dendrite only weakly responsive to this stimulus following learning. Note that during the learning process LTD could also occur in some synapses representing ‘yellow’ and ‘banana’ because these features were also part of the irrelevant stimuli (‘yellow strawberry’ and ‘red banana’, respectively). As a result, a small number of these synapses, in particular those whose initial weights are low, have been weakened and are not recruited in the clusters to encode the ‘yellow banana’ stimulus. This means that our learning rule tends to stabilize the number of synapses that are needed to perform the task, but not necessarily all the synapses carrying the relevant features (as some by chance may be outside of the plastic region of the LTP kernel). Depending on the initial local Ca response in the synapse to relevant and irrelevant stimuli, individual synapses might either be preferentially recruited into the LTP or LTD process.

Revised version: Figure 3C shows the evolution of synaptic conductances during the learning process for dendrite 1 (dendrite 2 is not shown, but has a similar behavior for relevant and irrelevant stimuli). The synapses representing the relevant feature combination in this dendrite (‘yellow’ and ‘banana’) are typically strengthened, eventually encoding this stimulus robustly. Conversely, the synapses for the feature ‘red’, activated during the irrelevant feature combination (‘red banana’) are all weakened, making the dendrite only weakly responsive to this stimulus following learning. Note that during the learning process LTD could also occur in some synapses representing ‘yellow’ and ‘banana’ because these features were also part of the irrelevant stimuli (‘yellow strawberry’ and ‘red banana’, respectively). As a result, a small number of these synapses, in particular those whose initial weights are low, have been weakened and are not recruited in the clusters to encode the ‘yellow banana’ stimulus. This means that our learning rule tends to stabilize the number of synapses that are needed to perform the task, but not necessarily all the synapses carrying the relevant features (as some by chance may be outside of the plastic region of the LTP kernel). Depending on the initial local calcium response in the synapse to relevant and irrelevant stimuli, individual synapses might either be preferentially recruited into the LTP or LTD process. An animated rendering of the first 400 training epochs for this panel, illustrating the frame-by-frame evolution of conductance and calcium in the example with clustered synapses, is provided in Figure 3—video 1.

Earlier version: The evolution of synaptic weights for clustered and feature-unspecific synapses is illustrated in the top panels of Figure 3C and D. Correspondingly, the calcium dynamics for three color-coded synapses in each case (indicated with arrows in the top panels) are shown in the bottom panels, demonstrating how the dynamics evolve under the influence of LTP and metaplasticity kernels. This highlights the process by which synapses stabilize at specific conductances/weights during learning.

Revised version: The evolution of synaptic weights for clustered and feature-unspecific synapses is illustrated in the top panels of Figure 3C and D. Correspondingly, the calcium dynamics for three color-coded synapses in each case (indicated with arrows in the top panels) are shown in the bottom panels, demonstrating how the dynamics evolve under the influence of LTP, LTD and metaplasticity. This highlights the process by which synapses stabilize at specific conductances/weights during learning.

Earlier version: The red synapses (exemplified in Figure 3C, second row) represents an irrelevant feature combination (‘red banana’), and as learning progresses, LTD becomes more frequent due to repeated dopamine pauses. These pauses shift the metaplasticity kernel towards higher calcium levels, moving the red synapses outside the range for LTP. This shift ensures that the red synapses weaken over time, preventing them from encoding irrelevant stimuli and reducing the dendrite’s response to these inputs. This can be compared with Figure 3—figure supplement 1 (red synapse) that lacks metaplasticity.

Revised version: The red synapses (exemplified in Figure 3C, second row) represent an irrelevant feature combination (‘red banana’), and as learning progresses, LTD becomes more frequent due to repeated dopamine pauses. These pauses shift the dynamic LTP and metaplasticity kernels towards higher calcium levels, moving the red synapses outside the range for LTP. This shift ensures that the red synapses weaken over time, preventing them from encoding irrelevant stimuli and reducing the dendrite’s response to these inputs. In setups lacking metaplasticity this does not happen, leading to an inability of the neuron to separate the response to relevant and irrelevant feature combinations (Figure 3—figure Supplement 1, red synapse).

Earlier version: The yellow synapse (exemplified in Figure 3C, third row), part of the relevant feature combination (‘yellow banana’), initially undergoes LTP as repeated dopamine rewards strengthen it. However, not all yellow synapses are equally strengthened. Synapses with weak initial weights may fail to potentiate further due to the slow adaptation of the LTP kernel and their inability to reach the required calcium levels. Specifically, the first example, “yellow syn. 1”, exhibits an increase in synaptic conductance, whereas the second example, “yellow syn. 2”, does not undergo potentiation (see arrows in Figure 3C top).

Revised version: Most synapses representing the relevant feature combination (‘yellow banana’, exemplified in Figure 3C, bottom rows) initially undergo LTP as repeated dopamine rewards strengthen them. However, synapses with weak initial weights may fail to potentiate further due to the slow adaptation of the LTP kernel and their inability to reach the required calcium levels. Specifically, for the synapses representing yellow, “Yellow syn 1” exhibits an increase in synaptic conductance, whereas “Yellow syn 2” does not undergo potentiation (see arrows in Figure 3C, top).

Removed in the revised version: The banana synapses (exemplified in Figure 3C, fourth row) follows a similar trajectory to the yellow synapses. As part of the relevant feature combination, they are strengthened during the learning process through repeated LTP. Some weak synapses also undergo LTD for the same reasons as for the few yellow synapses.

Earlier version: Feature-unspecific synapses, like the purple synapse in Figure 3D, typically weaken over time. These synapses are co-activated with irrelevant stimuli and experience frequent dopamine pauses, causing the metaplasticity kernel to shift upward and away from the calcium levels required for LTP. As a result, LTD dominates in these synapses, leading to their gradual weakening. Initially, LTP and LTD typically occur equally often, but as learning progresses, both the frequency of dopamine pauses and thus LTD decreases, allowing the weights to stabilize over time.

Revised version: Feature-unspecific synapses, like Example 1 and 3 in Figure 3D, typically weaken over time. These synapses are co-activated with irrelevant stimuli and experience frequent dopamine pauses, causing the metaplasticity kernel to shift upward and away from the calcium levels required for LTP. As a result, LTD dominates in these synapses, leading to their gradual weakening. Initially, LTP and LTD typically occur equally often, but as learning progresses, both the frequency of dopamine pauses and thus LTD decrease, allowing the weights to stabilize over time.

Earlier version: In contrast, the green synapse, though feature-unspecific, behaves differently due to its proximity to a cluster. It benefits from high local voltage and calcium, due to glutamate spillover that generates plateau potentials during relevant stimuli (e.g., ‘yellow banana’), keeping calcium levels within the high end of the LTP kernel. As the learning process progresses and dopamine pauses become less frequent, the LTP kernel shifts downward. This causes the calcium level associated with the other relevant stimuli (i.e. ‘red strawberry’) to also fall within the plastic region of the LTP kernel, and thereby further strengthening the synapse.

Revised version: In contrast, Example 2, though feature-unspecific, behaves differently due to its proximity to a cluster. It benefits from high local voltage and calcium during relevant stimuli (e.g. ‘yellow banana’), keeping calcium levels within the high end of the LTP kernel. As LTP dominates in this synapse, the LTP kernel gradually shifts down towards lower calcium concentrations. Eventually, this causes the calcium level associated with the other relevant stimulus (‘red strawberry’) to also fall within the plastic region of the LTP kernel, and thereby further strengthening the synapse.

Earlier version: Finally, the blue synapse initially triggers calcium levels within the LTP kernel, leading to early strengthening. However, as the LTD amplitude is proportional to the L-type calcium level, LTD eventually dominates, weakening the blue synapse. This happens because feature-unspecific synapses experience dopamine pauses and peaks equally often in the beginning. Once calcium levels fall below the optimal range of the LTP kernel due to dopamine pauses, LTP ceases, and the synapse stabilizes at a lower weight (compare to Figure 3—figure supplement 1, blue trace, where LTP and LTD stays balanced throughout the simulation and the weight stays high). In summary, synaptic weights are dynamically regulated by the interplay of LTD and LTP, modulated by the metaplasticity kernel. This process ensures that synapses adjust appropriately based on their relevance to feature combinations, stabilizing in a manner that reflects their functional roles and initial conditions.

Revised version: In summary, synaptic weights are dynamically regulated by the interplay of LTD and LTP, modulated by the metaplasticity kernel. This process ensures that synapses adjust appropriately based on their relevance to feature combinations, stabilizing in a manner that reflects their functional roles and initial conditions.

Added in the revised version: Clustering enables excitatory-only learning, but overlapping features make performance less reliable

Removed in the revised version: Characterization of input combinations that can be learned

Earlier version: After demonstrating that the SPN can differentiate between relevant and irrelevant stimuli in the simplified example in Figure 3, we generalized this setup by varying the innervation of the four features across the two dendrites and recorded the SPN’s performance on the NFBP as learning progressed. We only used feature innervations where both relevant feature combinations were present, with at least one relevant feature on each dendrite. This ensures sufficient innervation to potentially solve the NFBP (for an illustration see Figure 4A). The SPN’s performance was assessed during training to determine whether it spiked for relevant feature combinations and remained silent for irrelevant ones. Performance of 100% indicates the SPN spikes exclusively for relevant stimuli, while 50% indicates one of two scenarios: (i) the SPN spikes for all four stimuli, or (ii) it remains silent for all four. A performance of 75% reflects intermediate spiking behavior, such as spiking for one relevant combination while remaining silent for the other three, or spiking for three combinations and being silent for one irrelevant combination. The NFBP is considered solved when performance exceeds 87.5%, which occurs when the SPN consistently spikes for one relevant combination and spikes at least half the time for the other relevant combination while ignoring irrelevant stimuli.

Revised version: After demonstrating that the SPN can differentiate between relevant and irrelevant stimuli in the simplified example in Figure 3, we generalized this setup by varying the innervation of the four features across the two dendrites and recorded the SPN’s performance on the NFBP as learning progressed. We only used feature innervations where both relevant feature combinations were present, with at least one relevant feature on each dendrite. This ensures sufficient innervation to potentially solve the NFBP (for an illustration see Figure 4A). The SPN’s performance was assessed during training to determine whether it spiked for relevant feature combinations and remained silent for irrelevant ones. Performance of 100% indicates the SPN spikes exclusively for relevant stimuli, while 50% typically indicates one of two scenarios: (i) the SPN spikes for all four stimuli, or it remains silent for all four. A performance of 75% reflects intermediate spiking behavior, such as spiking for one relevant combination while remaining silent for the other three, or spiking for three combinations and being silent for one irrelevant combination. The NFBP is considered solved when performance exceeds 87.5%, which for example occurs when the SPN consistently spikes for one relevant combination and spikes at least half the time for the other relevant combination while ignoring irrelevant stimuli.

Impact of input configurations and synaptic cluster locations on NFBP learning performance

Removed in the revised version: (A): Illustration of creating different input configurations where two, three or four features are placed in two dendritic locations.

Earlier version: (B) : Performance trajectories for all 31 unique input configurations over training epochs. Each line represents the learning performance of a single configuration, with blue traces indicating setups where no dendrite receives more than three features and deep red traces representing setups where at least one dendrite contains all four features. The inset shows the number of configurations in each group that reached the NFBP learning criterion: 18/18 for the blue group and 5/13 for the deep red group.

Revised version: (A): Illustration of creating different input configurations where two, three (light blue circles), or four features (deep red circles) are placed in two dendritic locations. (B) : Performance trajectories for all 31 unique input configurations over training epochs. Each colored line represents the learning performance of a single configuration, with blue traces indicating setups where no dendrite receives more than three features and red traces representing setups where at least one dendrite contains all four features. The black line and grey shaded area represent the mean and standard deviation over all trials. The inset shows the number of configurations in each group that reached the NFBP learning criterion: 18/18 for the blue group and 5/13 for the deep red group. (C) : Performance on the last 160 stimuli for the two groups of configurations: (C1) setups where each dendrite has at most three features and (C2) setups where at least one dendrite contains all four features. (D) and (E): Performance in a three-feature configuration as a function of cluster location. (D) illustrates the distribution of synaptic clusters across dendritic locations, and (E) shows performance over the last 160 stimulus presentations as a function of somatic distance for these configurations. A total of 60 unique configurations were tested, with synapse clusters randomly assigned to different dendritic locations. The solid line represents a quadratic fit to the data.

Removed in the revised version: (C) : Performance on the last 160 stimuli for the two groups of configurations: (C1) setups where each dendrite has at most three features and (C2) setups where at least one dendrite contains all four features.

Removed in the revised version: (D) and (E): Performance in a three-feature configuration as a function of cluster location. (D) illustrates the distribution of synaptic clusters across dendritic locations, and (E) shows performance over the last 160 stimulus presentations as a function of somatic distance for these configurations. A total of 60 unique configurations were tested, with synapse clusters randomly assigned to different dendritic locations. The solid line represents a quadratic fit to the data.

Earlier version: In cases where both dendrites receive two or three features, ensuring that each dendrite has at least one relevant combination, the mean performance exceeds 90%, indicating that both relevant combinations are learned (Figure 4B, blue traces). However, when at least one dendrite is innervated by all four features, learning becomes significantly more challenging. The dendrite with four features struggles to resolve the competition between feature combinations, making it difficult to encode a single relevant combination. This often results in only one relevant combination being learned, with the same combination often being redundantly encoded in both dendrites, or else the dendrite with four features is failing to encode any combination as all features are weakened (Figure 4B, deep red traces). In some cases, depending on the order of stimuli during training, both relevant combinations can still be learned, solving the NFBP. Specifically, 5 out of 13 setups in which one dendrite received all four features successfully solved the task (Figure 4B).

Revised version: In cases where both dendrites receive two or three features, ensuring that each dendrite has at least one relevant combination, the mean performance exceeds 90%, indicating that both relevant combinations are learned (Figure 4B, blue traces). However, when at least one dendrite is innervated by all four features, learning becomes significantly more challenging. The dendrite with four features struggles to resolve the competition between feature combinations, making it difficult to encode a single relevant combination (Figure 4B, deep red traces). However, in some cases, depending on the order of stimuli during training, both relevant combinations can still be learned, solving the NFBP. Here, 5 out of 13 setups in which one dendrite received all four features successfully solved the task (Figure 4B, inset).

Earlier version: Figure 4C provides a detailed view of performance during the last 160 training examples for each of the four stimuli. Figure 4C1 corresponds to the blue traces in Figure 4B, where the NFBP is solved, with most errors caused by spiking for irrelevant combinations. Figure 4C2 corresponds to the deep red traces in Figure 4B, where the NFBP is generally not solved, as relevant stimuli elicit spikes only about 75% of the time. These results suggest that when one dendrite is innervated by all four features, the competition between feature combinations creates a bottleneck for learning. This can lead to one combination dominating early in training and being encoded in both dendrites, while the other relevant combination is weakened or entirely suppressed. In cases where learning fails, the dendrite with four features often either fails to encode any combination due to excessive LTD or instead encodes the same combination already represented in the other dendrite. Addressing this challenge may require mechanisms to “break symmetry” during learning, as noted in previous studies using abstract neuron models. Symmetry-breaking strategies, such as amplifying or attenuating dendritic nonlinearities as they propagate to the soma, can reduce competition and enable successful differentiation of feature combinations (e.g., branch plasticity; Legenstein and Maas, 2011).

Revised version: Figure 4C provides a detailed view of performance during the last 160 training examples for each of the four stimuli. Figure 4C1 corresponds to the blue traces in Figure 4B, where the NFBP is considered solved. Figure 4C2 corresponds to the deep red traces in Figure 4B, where the NFBP is generally not solved, as relevant stimuli elicit spikes only about 75% of the time. These results suggest that when one dendrite is innervated by all four features, the competition between feature combinations creates a bottleneck for learning. In cases where learning fails, the dendrite with four features often either fails to encode any combination due to excessive LTD or instead encodes the same combination already represented in the other dendrite. Addressing this challenge may require mechanisms to “break symmetry” during learning, as noted in previous studies using abstract neuron models. Symmetry-breaking strategies, such as amplifying or attenuating dendritic nonlinearities as they propagate to the soma, can reduce competition and enable successful differentiation of feature combinations (e.g., branch plasticity; Legenstein and Maass 2011).

Optimal learning is achieved at intermediate distances from soma through excitatory plasticity

Earlier version: We also investigated the impact of synaptic positioning on learning (Figure 4D) when using the same input setup as in Figure 3, but varying dendritic locations. Our results predict that the best performance on the NFBP is obtained with synaptic clusters positioned at intermediate somatic distances from the soma (Figure 4E). From Figure 4E one can infer that after learning the proximal synapses have actually decreased when trained on the NFBP as the performance is around 50%. The neuron stays silent for the irrelevant feature combinations (blue dots), which is the correct response, but also the neuron can’t respond to the relevant feature combinations (red dots). For successively more distal synapses the performance increases and then slightly decreases for the most distal clusters that sometimes fail to evoke somatic spiking for the correct feature combinations (red dots). This result can be conceptually explained in the following way. The electrotonic properties of dendrites dictate that synapses near the soma, in the most proximal regions, are less capable of supralinear input integration underlying plateau potentials (Du et al., 2017). This is due to the soma acting as a current sink, resulting in smaller localized voltage changes (and hence a lower input resistance in accordance with Ohm’s law). Consequently, these synapses cannot easily evoke dendritic nonlinearities necessary for solving the NFBP, and hence the performance with proximal clusters is low. Note that in our simulations we allow glutamatergic synapses on spines quite close to the soma, although SPN dendritic spines are relatively rare at more proximal distances than 40-50 μm from the soma (Wilson et al., 1983).

Revised version: We also investigated the impact of synaptic positioning on learning (Figure 4D) when using the same input setup as in Figure 3, but varying dendritic locations. Our results predict that the best performance on the NFBP is obtained with synaptic clusters positioned at intermediate somatic distances from the soma (Figure 4E). From Figure 4E one can infer that after learning the proximal synapses have all been decreased to such an extent that the neuron has stopped spiking for all four combinations (as the total performance is around 50% and relevant feature combinations is at 0%). For successively more distal synapses the performance increases and then slightly decreases for the most distal clusters that sometimes fail to evoke somatic spiking for the correct feature combinations (red dots). This result can be conceptually explained in the following way. The electrotonic properties of dendrites dictate that synapses near the soma, in the most proximal regions, are less capable of supralinear input integration underlying plateau potentials (Du et al., 2017). This is due to the soma acting as a current sink, resulting in smaller localized voltage changes (and hence a lower input resistance in accordance with Ohm’s law). Consequently, these synapses cannot easily evoke dendritic nonlinearities necessary for solving the NFBP, and hence the performance with proximal clusters is low. Note that in our simulations we allow glutamatergic synapses on spines quite close to the soma, although SPN dendritic spines are relatively rare at more proximal distances than 40–50 µm from the soma (Wilson et al., 1983).

In contrast, the most distal dendritic regions are electrically more isolated and have a higher local input resistance, enabling larger voltage changes locally and thus also higher local calcium concentrations when synapses are activated in our model. This allows even a small number of active synapses to generate local supralinear NMDA-dependent responses. Such ease of elevating the local calcium, seems advantageous, however, this ultimately results in decreased performance on the NFBP for the following reasons. In distal synaptic clusters, excessive spiking for irrelevant stimuli occurs during the early and middle stages of learning. This leads to frequent punishment signals (negative feedback), which, in turn, reduce the neuron’s ability to spike for relevant combinations later in training. Additionally, while distal synapses can generate plateau potentials and NMDA-dependent responses, these signals naturally attenuate as they propagate toward the soma, sometimes failing to elicit somatic spikes for the relevant stimuli. This combination of excessive negative feedback early on and attenuated distal signals reduces the effectiveness of the feedback loop, making it difficult for the neuron to selectively strengthen synapses associated with relevant stimuli and weaken those tied to irrelevant ones.

Earlier version: The ideal learning zone for NFBP in our simulations thus lies at an intermediate somatic distance, roughly 100-150 µm from the soma, where synapses can effectively contribute to learning the NFBP (Figure 4E). In this zone synaptic changes are more likely to impact the neuron’s firing probability as the dendritic plateau potential at this location causes a larger elevation of the somatic potential, and thus synapses at this distance benefit more from the dopamine feedback loop. Note that the prediction that proximal and very distal synapses are less likely to contribute to the solving of the NFBP doesn’t imply they are not important for more ‘linear’ learning contexts. For instance, if we had trained the neuron to only respond to one single stimulus, such as ‘red strawberry’, both proximal and very distal synapses representing that stimulus would of course be able to both strengthen or weaken as well as contribute to spiking of the neuron following learning.

Revised version: In our simulations, the ideal learning zone for NFBP thus lies at an intermediate somatic distance, roughly 100–150 µm from the soma, where synapses can effectively contribute to learning the NFBP (Figure 4E). In this zone synaptic changes are more likely to impact the neuron’s firing probability as the dendritic plateau potential at this location causes a larger elevation of the somatic potential, and thus synapses at this distance benefit more from the dopamine feedback loop. Note that the prediction that proximal and very distal synapses are less likely to contribute to the solving of the NFBP doesn’t imply they are not important for more ‘linear’ learning contexts. For instance, if we had trained the neuron to only respond to one single stimulus, such as ‘red strawberry’, both proximal and very distal synapses representing that stimulus would of course be able to both strengthen or weaken as well as contribute to spiking of the neuron following learning.

Added in the revised version: Inhibitory plasticity can enhance contrast and robustness of dendritic computations

Removed in the revised version: The possible role of inhibitory plasticity in learning

In our initial simulations we assumed that only the excitatory synapses could undergo plasticity during learning, and we identified two critical observations that highlight possible areas for improvement. The first is a vulnerability to noise, which resulted in a performance around 90 percent, as illustrated in Figure 4B. The second is a decrease in performance observed across very distal synapses, as detailed in Figure 4E. These findings prompted us to explore the potential role of inhibitory synapses.

Earlier version: We developed a phenomenological inhibitory plasticity rule to enhance the contrast of dendritic nonlinearities (detailed in the Methods section). This rule is designed to compartmentalize dendrites, ensuring they predominantly respond to excitatory inputs that cause the strongest activation. Such compartmentalization has been experimentally observed in radial oblique and basal dendrites of CA1 pyramidal neurons, where NMDA spikes and plateau potentials are the main forms of dendritic nonlinearity. Co-located inhibitory synapses in these dendrites regulate whether these nonlinear responses are elicited or not (Lovett-Barron et al., 2012).

Revised version: We developed a phenomenological inhibitory plasticity rule to enhance the contrast of dendritic nonlinearities (see Figure 5A, detailed in the Methods section). This rule is designed to compartmentalize dendrites, ensuring they predominantly respond to excitatory inputs that cause the strongest activation. Such compartmentalization has been experimentally observed in radial oblique and basal dendrites of CA1 pyramidal neurons, where NMDA spikes and plateau potentials are the main forms of dendritic nonlinearity. Co-located inhibitory synapses in these dendrites regulate whether these nonlinear responses are elicited (Lovett-Barron et al., 2012; Milstein et al., 2015; Grienberger et al., 2017).

Removed in the revised version: Unlike excitatory plasticity, which relies on dopaminergic feedback signals, our inhibitory plasticity model follows a rule that passively follows the local Ca level from voltage-gated channels. It reinforces the voltage elevation of the most active excitatory synapses within a dendritic branch by decreasing the inhibition corresponding to the same features there (Chapman, Nuwer and Jacob, 2022), and conversely, it strengthens inhibitory synapses for features which generate less excitatory activity (see Methods for details).

Removed in the revised version: To demonstrate how the inhibitory plasticity rule works, we use the same excitatory synapse setup as in Figure 4 to which we add four inhibitory synapses near each cluster in the middle of the dendritic branch, representing each of the four features (Figure 5A). Thus, a single feature activates both the excitatory and inhibitory synapses. To achieve a level of depolarization and spike probability comparable to that in our excitatory-only setup, we increased the number of feature-unspecific input from 108 to 144. Alongside this, we began with low inhibitory synaptic weights. These two changes were needed in order to maintain higher baseline activity in the model as starting with strong inhibitory weights could excessively suppress excitatory activity, as e. g. inhibitory inputs close to clustered synapses can effectively counteract the NMDA-dependent nonlinearities (Doron et al., 2017; Du et al., 2017; Dorman, Jędrzejewska-Szmek and Blackwell, 2018). Here, as in the example in Figure 3, the conductances of the excitatory inputs representing ‘yellow’ and ‘banana’ increase in dendrite 1, strengthening the ‘yellow banana’ pairing, while the weights of synapses representing the feature ‘red’ decrease (Figure 5E). Conversely, the inhibitory synapses associated with the ‘yellow’ and ‘banana’ features in dendrite 1 are weakened, while those linked to the ‘red’ and ‘strawberry’ features in the same dendrite are strengthened (Figure 5D, right panel). This behaviour of the inhibitory plasticity rule effectively prevents the cell from spiking following activation of the irrelevant stimuli. In this particular example the ‘red banana’ combination is strongly inhibited and thereby effectively compartmentalizing dendrite 1 to be responsive to ‘yellow banana’ (since in this case no excitatory ‘strawberry’ input is present).

Earlier version: Effects of inhibitory inputs on performance

Revised version: Effects of inhibitory plasticity on performance

Added in the revised version: (A): Dendritic input configuration with inhibitory synapses added. The setup of the excitatory and feature-unspecific synapses is the same as in Figure 4A. Plastic inhibitory synaptic connections for each of the four features are added in both dendrites, with one synapse per feature. (B) : Displays average performance for 31 unique input configurations of 2, 3, or 4 features on two dendrites as a comparison between the setup with (orange) and without inhibitory plasticity (blue). Shaded areas show standard deviation. The bar plot inset shows the number of configurations, with 4 features in at least one dendrite, that successfully solved the NFBP with and without inhibitory plasticity. (C) : Performance over the last 160 stimulus presentations as a function of somatic distance of the synaptic clusters for the input configuration in Figure 4D, with added inhibition. A total of 60 unique configurations were tested, with synapse clusters randomly assigned to different dendritic locations. The solid line represents a quadratic fit to the data. The dashed line is the corresponding quadratic fit from Figure 4E. (D) : Peak voltage-gated calcium (left panel, dots) and plasticity threshold dynamics (lines), in dendrite 1, over training epochs. The upper threshold (θinh,high) moves towards peak calcium levels while the lower threshold (θinh,low) moves towards the next highest level. (Right) Inhibitory synaptic conductances for the synapses in dendrite 1. Strengthened inhibitory synapses prevent excitation by the corresponding excitatory features (here ‘red’ and ‘strawberry’), while weakened inhibitory synapses allow excitation by their corresponding features (here ‘yellow’ and ‘banana’). (E) : Excitatory synaptic conductances (top) and calcium levels and kernel dynamics (below) over learning. The specific conductances in the top panel corresponding to the kernel dynamics and calcium levels in the bottom panels are indicated with arrows. Note that kernel and calcium dynamics for the example ‘yellow’ synapse refer to the only ‘yellow’ synapse in the top panel which initially weakens and is later strengthened. The solid line shows the midpoint of the LTP kernels. Max refers to the peak NMDA calcium during a single stimulus. The darker shaded regions represent the LTP kernel and the lighter shaded ones show the wider metaplasticity kernel.

Added in the revised version: Learning with and without inhibitory plasticity in distal dendrites.

Added in the revised version: (A) : The figure illustrates the impact of inhibitory plasticity on learning dynamics in two distal dendrites (d1 and d2). The top panels depict somatic voltage, while the middle and bottom panels show dendritic voltages. The left panels show voltage in the middle of the learning process (approximately 480 training epochs) with and without inhibitory plasticity. The right panels show voltage at the end of the learning process (approximately 960 training epochs). The dashed black line corresponds to the scenario without inhibitory plasticity, and the solid red line indicates the presence of inhibitory plasticity. In the middle of learning, inhibition prevents spiking for the irrelevant stimuli, thus leading to less LTD. This preserves synaptic strengths in both the clustered and feature-unspecific synapses, allowing the neuron to spike for the relevant stimuli at the end of learning. The overall decrease in unspecific input was 39% without inhibitory plasticity and 35% with inhibitory plasticity. (B) : The effect of inhibitory plasticity on learning in a dendrite receiving all four features. The left panel shows excitatory synaptic weights in d1 without inhibitory plasticity, where all features weaken over time due to competition. The right panel illustrates excitatory synapses in the presence of inhibitory plasticity, where learning is better regulated.

Added in the revised version: Unlike excitatory plasticity, which relies on dopaminergic feedback signals, our inhibitory plasticity model follows a rule that passively follows the local calcium level from voltage-gated channels. It reinforces the voltage elevation of the most active excitatory synapses within a dendritic branch by decreasing the inhibition corresponding to the same features there (Chapman et al., 2022), and conversely, it strengthens inhibitory synapses for features which generate less excitatory activity (see Figure 5A, right panel, illustrating how the rule updates active and inactive inhibitory synapses).

Added in the revised version: To demonstrate how the inhibitory plasticity rule works, we use the same excitatory synapse setup as in Figure 4 to which we add four inhibitory synapses near each cluster in the middle of the dendritic branch, representing each of the four features (Figure 5A). Thus, a single feature activates both the excitatory and inhibitory synapses. To achieve a level of depolarization and spike probability comparable to that in our excitatory-only setup, we increased the number of featureunspecific inputs from 108 to 144. Alongside this, we began with low inhibitory synaptic weights. These two changes were needed in order to maintain higher baseline activity in the model, as starting with strong inhibitory weights could excessively suppress excitatory activity, since inhibitory inputs close to clustered synapses can effectively counteract the NMDA-dependent nonlinearities (Doron et al., 2017; Du et al., 2017; Dorman et al., 2018). Here, as in the example in Figure 3, the conductances of the excitatory inputs representing ‘yellow’ and ‘banana’ increase in dendrite 1, strengthening the ‘yellow banana’ pairing, while the weights of synapses representing the feature ‘red’ decrease (Figure 5E). Conversely, the inhibitory synapses associated with the ‘yellow’ and ‘banana’ features in dendrite 1 are weakened, while those linked to the ‘red’ and ‘strawberry’ features in the same dendrite are strengthened (Figure 5D, right panel). This behaviour of the inhibitory plasticity rule effectively prevents the cell from spiking following activation of the irrelevant stimuli. In this particular example the ‘red banana’ combination is strongly inhibited, thereby effectively compartmentalizing dendrite 1 to be responsive to ‘yellow banana’ (since in this case no excitatory ‘strawberry’ input is present).

Removed in the revised version: (A) Dendritic input configuration with inhibitory synapses added. The setup of the excitatory and feature-unspecific synapses is the same as in Figure 3A (the feature-unspecific synapses are not shown). Inhibitory synaptic connections for each of the four features are added in both dendrites, with one synapse per feature.

Removed in the revised version: (B) : Displays average performance for 31 unique input configurations of 2, 3, or 4 features on two dendrites as a comparison between the setup with (orange) and without inhibitory plasticity (blue). The bar plot inset shows the number of configurations that successfully solved the NFBP with and without inhibition.

Removed in the revised version: (C) : Performance over the last 160 stimulus presentations as a function of somatic distance of the synaptic clusters for the input configuration in (A). A total of 60 unique configurations were tested, with synapse clusters randomly assigned to different dendritic locations. The solid line represents a quadratic fit to the data. The dashed line is the corresponding quadratic fit from Figure 4E.

Removed in the revised version: (D) : Peak voltage-gated calcium (left panel, dots) and plasticity thresholds dynamics (lines) during learning (number of training examples). The upper threshold (CTH) moves towards peak calcium levels while the lower threshold (CTL) moves towards the next highest level. (Right) Inhibitory synaptic conductances for the synapses in dendrite 1. Strengthened inhibitory synapses prevent excitation by the corresponding features (here ‘red’ and ‘strawberry’), while weakened inhibitory synapses allow excitation by their corresponding features (here ‘yellow’ and ‘banana’).

Removed in the revised version: (E) : Excitatory synaptic conductances (top) and calcium levels and kernel dynamics (below) over learning. The specific conductances in the top panel corresponding to the kernel dynamics and calcium levels in the bottom panels are indicated with arrows. Note that kernel and calcium dynamics for the example ‘yellow’ synapse refer to the only ‘yellow’ synapse in the top panel which initially weakens and is later strengthened. MP-K, LTP-K stands for metaplasticity- and LTP kernel, while the solid line shows the midpoint of the LTP kernels. Max refers to the peak NMDA calcium during a single stimulus.

Earlier version: We also show the dynamics of peak calcium levels associated with both excitatory and inhibitory synapses for each task (Figure 5D, left panel, and Figure 5E). For excitatory synapses, the patterns in peak NMDA calcium levels also behave as in the example in Figure 3C (compared to Figure 5E). For inhibitory synapses, the peak calcium levels arising from voltage-gated calcium channels further exemplify how inhibition suppresses calcium activity arising from irrelevant stimuli while allowing that evoked by relevant stimuli (Figure 5D, left panel). The maximum threshold level (CTH) of calcium for inhibitory synapses followed the highest excitatory synaptic activity, which in this case corresponded to the ‘yellow banana’ input.

Revised version: We also show the dynamics of peak calcium levels associated with both excitatory and inhibitory synapses for each task (Figure 5D, left panel, and Figure 5E). For excitatory synapses, the patterns in peak NMDA calcium levels behave as in the example in Figure 3C (compare to Figure 5E). For inhibitory synapses, the peak calcium levels arising from voltage-gated calcium channels further exemplify how inhibition suppresses calcium activity arising from irrelevant stimuli while allowing that evoked by relevant stimuli (Figure 5D, left panel). The high threshold level of calcium for inhibitory synapses (θinh, high) followed the highest excitatory synaptic activity, which in this case corresponded to the ‘yellow banana’ input.

Earlier version: We also compared the performance for different feature configurations with added inhibitory synapses to the results for the excitatory-only setup from Figure 4B. The results show not only that learning with inhibitory plasticity is faster, i.e. requires less training examples, but also achieves high performance, nearing 100%. In the configurations with all four features in one dendrite 7/13 learned the NFBP, compared to 5/13 without inhibition (shown in the bar plot inset in Figure 5B).

Revised version: We also compared the performance for different feature configurations with added inhibitory synapses to the results for the excitatory-only setup from Figure 4B. The results show not only that learning with inhibitory plasticity is faster, i.e. requires fewer training examples, but also achieves high performance, nearing 100%. In the configurations with all four features in one dendrite 7/13 learned the NFBP, compared to 5/13 without inhibition (shown in the bar plot inset in Figure 5B).

Earlier version: To understand why inhibition improved learning, we examined the dynamics of L-type calcium during early learning. The presence of inhibition resulted in slightly lower initial L-type calcium levels compared to the excitatory-only setup, leading to reduced LTD at the start. This reduction in LTD helped prevent excessive weakening of synapses, increasing the likelihood of selecting one combination over others (Figure 5—figure supplement 1B, showing excitatory synaptic weights without (left) and with inhibition (middle): In the setup without inhibition all four features weaken over time due to competition, while in contrast, the setup including inhibition demonstrates better-regulated learning and more stable feature selection).

Revised version: To understand why inhibitory plasticity improved learning, we examined the dynamics of L-type calcium during early learning. The presence of inhibition resulted in slightly lower initial L-type calcium levels compared to the excitatory-only setup, leading to reduced LTD at the start. This reduction in LTD helped prevent excessive weakening of synapses, increasing the likelihood of selecting one combination over others (Figure 5—figure Supplement 1B, showing excitatory synaptic weights without (left) and with inhibitory plasticity (middle): in the setup without inhibitory plasticity all four features weaken over time due to competition, while in contrast, the setup including inhibitory plasticity demonstrates better-regulated learning and more stable feature selection).

Inhibitory plasticity additionally produced a marked improvement in performance when the location of the synaptic clusters was varied, especially prominent at very distal locations (Figure 5C). Without inhibition it was difficult for synapses on distal dendrites to differentiate between relevant and irrelevant stimuli. In the middle of the learning, setups without inhibition typically spiked too much compared to setups including inhibition (e.g. spike for ‘red banana’ in the left panels of Figure 5—figure Supplement 1A). This in turn caused increased LTD and decreased weights in the feature-unspecific synapses, leading to too little spiking at the end of the learning (e.g. no spike for ‘yellow banana’ in the right panels of Figure 5—figure Supplement 1A). With inhibition, on the other hand, the contrast between relevant and irrelevant stimuli was larger. Larger contrast in turn gave less spiking for the wrong combination in the middle of learning and thereby less reduction of the weights of the feature-unspecific synapses. Therefore, large excitation in a single dendrite, combined with a larger drive of feature-unspecific synapses were enough to cause spiking in relevant stimuli at the end of learning, while spiking in irrelevant stimuli was suppressed (Figure 5—figure Supplement 1A, red traces).

In summary, incorporating inhibitory plasticity suggests that inhibitory synapses may fine-tune dendritic responsiveness and enhance NFBP performance by preventing excitation from irrelevant stimuli, which in turn allows for more rapid, robust, and accurate learning.

Added in the revised version: Inhibitory plasticity improves performance with randomly distributed inputs

Removed in the revised version: Dendritic Branch Plasticity and Synaptic Learning – Investigating the Role of Randomly Distributing Excitatory and Inhibitory Synapses on the Dendrites

Earlier version: We finally challenged our plasticity rule by relaxing the assumption that single features are represented by pre-clustered synapses on specific dendritic branches. The behaviour of the synaptic plasticity rule was therefore investigated using a setup where 200 excitatory synapses, representing the different features, were randomly distributed across 30 dendrites. In this setup, each feature was represented by 40 synapses and an additional 40 represented the feature-unspecific synapses. Figure 6A illustrates the setup and exemplifies the pre- and post-learning synaptic weights for both excitatory and inhibitory synapses. Our objective was to examine the learning dynamics in the absence of assumed synaptic clustering and to determine the capability of the single neuron to learn the NFBP. To address the reduced effectiveness of non-clustered synaptic inputs in eliciting sufficient depolarization and calcium influx—key for learning and synaptic plasticity—we slightly increased the initial weights compared to the clustered setup to 0.3 ± 0.1 (around 0.75 nS). A branch-specific mechanism for creating non-linearities was also introduced by assuming that glutamate spillover occurs in all co-activated synapses on the same dendritic branch. This hypothetical construct assumes that concurrent activation of a critical mass of co-activated synapses in a single dendritic branch would trigger both plateau potentials and a rise in calcium levels, endowing the dendrites with supralinear responses even without closely clustered synapses. Starting with an excitation-only setup, and running the model with these changes resulted in a moderate mean performance of about 65% (Figure 6B, ‘Without inhibitory’). However, on a linear task (e.g. learning only ‘yellow banana’ or ‘red strawberry’) with randomly placed synapses the performance was close to 100% (even when starting from an initial condition without spiking—a paradigm called subthreshold learning—see Figure 6—figure supplement 1B and 1C).

Revised version: We finally challenged our plasticity rule by relaxing the assumption that single features are represented by pre-clustered synapses on specific dendritic branches. The behaviour of the synaptic plasticity rule was therefore investigated using a setup where 200 excitatory synapses, representing the different features, were randomly distributed across 30 dendrites. In this setup, each feature was represented by 40 synapses and an additional 40 represented the feature-unspecific synapses. Figure 6A illustrates the setup and exemplifies the preand post-learning synaptic weights for both excitatory and inhibitory synapses. Our objective was to examine the learning dynamics in the absence of assumed synaptic clustering and to determine the capability of the single neuron to learn the NFBP. To address the reduced effectiveness of non-clustered synaptic inputs in eliciting sufficient depolarization and calcium influx—key for learning and synaptic plasticity—we slightly increased the initial weights compared to the clustered setup to 0.3 ± 0.1 (around 0.75 nS). A branch-specific mechanism for creating non-linearities was also introduced by assuming that glutamate spillover occurs in all co-activated synapses on the same dendritic branch. This hypothetical construct assumes that concurrent activation of a critical mass of co-activated synapses in a single dendritic branch would trigger both plateau potentials and a rise in calcium levels, endowing the dendrites with supralinear responses even without closely clustered synapses.

Performance analysis of learning using distributed synaptic inputs

Earlier version: (A) : Example illustration of synaptic distribution before (top) and after learning (bottom) of the 200 excitatory and 60 inhibitory inputs. S - strawberry, Y - yellow, B - banana, R - red, U - feature-unspecific inputs. Bordered markers (circles and diamonds) indicate synapses whose conductances are shown in (C).

Revised version: (A) : Example illustration of synaptic distribution before (top) and after learning (bottom) of the 200 excitatory and 60 inhibitory inputs. S — strawberry, Y — yellow, B — banana, R — red, U — feature-unspecific inputs. Bordered markers (circles and diamonds) indicate synapses whose conductances are shown in (C). (B) : Performance over training epochs with and without inhibitory plasticity. (Left) Mean performance of the setups with and without inhibitory plasticity. Shaded areas show standard deviation. (Right) Individual traces for the setup with inhibitory plasticity. Each of these individual trials (31 in total) uses a unique random distribution of synapses. (C) : Example of summed synaptic conductances (left) and voltage (right) in the soma, and four example dendrites (d1–d4) of the synaptic distribution in (A) (corresponding to a trial where the NFBP is successfully solved). The sums of both excitatory (left) and inhibitory (right) inputs are shown.

Added in the revised version: Other learning paradigms.

Added in the revised version: (A) : Learning of 5×5 and 3×3 feature combination tasks with randomly distributed synapses over 30 dendrites. The neuron is trained to spike for black-marked combinations and remain silent for white-marked ones. For the 25-feature (5×5) combination, the model used 500 excitatory synapses (50 per feature) with initial conductances of 1.125 ± 0.25 nS, along with 100 inhibitory synapses (10 per feature) with initial conductances of 0.1±0.01 nS. For the 9-feature (3×3) combination, the model used 312 excitatory synapses (52 per feature) and 96 inhibitory synapses (16 per feature), with the same initial weights as in the 5×5 case. (B) : Learning of a 2×2 linear feature combination task with randomly distributed synapses. Each feature is represented by 75 synapses distributed across 30 dendrites (300 total synapses), with an initial synaptic weight of 0.375±0.25 nS. The top row shows the feature combinations used, followed by the performance across 30 different synaptic distributions. The bottom row illustrates somatic voltage traces before and after learning, where the neuron starts from a silent state (red trace) and learns to spike through metaplasticity and dopamine-driven weight updates. Additionally, in the subthreshold learning paradigm, dopamine was given without spiking for expected firings and suppressed when unexpected spiking occurred. (C) : Evolution of synaptic conductances (left column) and calcium dynamics with metaplasticity activity (right column) during learning for each feature (X1, X2, Y1, Y2). The left panels display the weight adjustments of all synapses corresponding to each feature, with one synapse per feature bolded to highlight its individual trajectory. The right panels show the corresponding calcium traces and kernel dynamics for the bolded synapse. The solid line represents the midpoint of the kernels where LTP is strongest. ‘Max’ indicates the peak NMDA calcium level during a single stimulus. The darker shaded regions represent the LTP kernel and the lighter shaded ones show the wider meta-plasticity kernel.

Added in the revised version: Starting with an excitation-only setup resulted in a moderate mean performance of about 65% (Figure 6B, ‘Without inhibitory plasticity’). However, on a linear task (e.g. learning only ‘yellow banana’ or ‘red strawberry’) with randomly placed synapses the performance was close to 100% (even in so-called subthreshold learning, starting from an initial condition without spiking; see Figure 6— figure Supplement 1B and C).

Removed in the revised version: (B) : Performance over training epochs with and without inhibitory plasticity. (Left) Mean performance of the setups with and without inhibitory synapses. (Right) Individual traces for the setup with inhibition. Each of these individual trials (31 in total) uses a unique random distribution of synapses.

Removed in the revised version: (C) : Example of summed synaptic conductances (left) and voltage (right) in the soma, and four example dendrites (d1-d4) of the synaptic distribution in (A) in a trial where the NFBP is successfully solved. The sums of both excitatory (left) and inhibitory (right) inputs are shown.

Earlier version: We next extended our investigation by including 60 inhibitory synapses, 15 for each feature, dispersed randomly over the 30 dendrites. Performance improved with the introduction of these additional nonlinearities, especially in combination with the branch-specific spillover mechanisms. In the right panel of Figure 6B, we show that with inhibition the model reaches a performance of about 75% for all distributions of synapses, while a subset of distributions (5/31) achieve performance close to 100%. There are also a few distributions that partially solve the NFBP, reaching a performance between 75% and 100% (4/31).

Revised version: We next extended our investigation by including 60 inhibitory synapses, 15 for each feature, dispersed randomly over the 30 dendrites. Performance improved with the introduction of these additional nonlinearities, especially in combination with the branch-specific spillover mechanisms. In the right panel of Figure 6B we show that, with inhibition, most individual distributions (22/31) reach a performance of about 75%, while a subset of distributions (5/31) achieve performance close to 100%. There are also a few distributions that partially solve the NFBP, reaching a performance between 75% and 100% (4/31).

Earlier version: To explore the model’s ability to handle increased feature complexity, we extended the task to include two configurations: a subset of 9 feature combinations, including one non-linear combination, and a more challenging subset of 25 feature combinations, including two non-linear combinations (see Figure 6—figure supplement 1A). In these extended simulations, the model achieved accuracies of 80% and 75% for the 9- and 25-feature tasks, respectively, as shown in see Figure 6—figure supplement 1A.

Revised version: To explore the model’s ability to handle increased feature complexity, we extended the task to include two configurations: a subset of 9 feature combinations, including one non-linear combination, and a more challenging subset of 25 feature combinations, including two non-linear combinations (see Figure 6—figure Supplement 1A). In these extended simulations, the model achieved accuracies of 80% and 75% for the 9- and 25-feature tasks, respectively, as shown in Figure 6— figure Supplement 1A.

Earlier version: For a more granular analysis of the successful cases with randomly distributed synapses, we selected an example that successfully learned the NFBP. In Figure 6C, we present the somatic and dendritic voltage responses for four dendritic branches (d1–d4) where relevant stimuli were encoded. The voltage traces display somatic and dendritic responses to all four stimuli after learning. Additionally, the cumulative synaptic conductances for both excitatory and inhibitory inputs were calculated at the midpoint of each dendritic branch. Notably, dendrites 1 and 4 show enhanced excitatory inputs for the ‘yellow banana’ pattern and increased inhibitory inputs for the ‘red strawberry’ pattern, while dendrites 2 and 3 exhibit the reverse arrangement. The synapses in dendrites 1 to 4, whose weights are shown in panel C, are marked in Figure 6A with black borders around the markers.

Revised version: For a more granular analysis of the successful cases with randomly distributed synapses and assumed branch-specific plasticity, we looked in detail at the conductances of an example that successfully learned the NFBP. In Figure 6C we present the somatic and dendritic voltage responses for four dendritic branches (d1–d4) where relevant stimuli were encoded. The voltage traces display somatic and dendritic responses to all four stimuli after learning. Additionally, the cumulative synaptic conductances for both excitatory and inhibitory inputs were calculated at the midpoint of each dendritic branch. Notably, dendrites 1 and 4 show enhanced excitatory inputs for the ‘yellow banana’ pattern and increased inhibitory inputs for the ‘red strawberry’ pattern, while dendrites 2 and 3 exhibit the reverse arrangement. The synapses in dendrites 1 to 4, whose weights are shown in panel C, are marked in Figure 6A with black borders around the markers.

Earlier version: These results show that for an adequate random innervation of distributed synapses, where the necessary features for a relevant stimulus innervate the same dendritic branch with enough synapses, that stimulus can be stored on that dendrite. In this way, the NFBP can also be learned if the two relevant stimuli are encoded on different dendrites and each of them can trigger a supralinear dendritic response. Since on average the two features representing a relevant stimulus do not innervate a single dendrite with enough synapses, the stimulus is not stored in a single dendrite, but is distributed across the dendritic tree. And since distributed synapses summate more linearly at the soma, only one of the relevant stimuli can be typically encoded by the neuron.

Revised version: These results show that for an adequate random innervation of distributed synapses—where the necessary features for a relevant stimulus innervate the same dendritic branch with enough synapses—that stimulus can be stored on that dendrite. In this way, the NFBP can also be learned if the two relevant stimuli are encoded on different dendrites and each of them can trigger a supralinear dendritic response. Since, on average, the two features representing a relevant stimulus do not innervate a single dendrite with enough synapses, the stimulus is not stored in a single dendrite but is distributed across the dendritic tree. And since distributed synapses summate more linearly at the soma, only one of the relevant stimuli can typically be encoded by the neuron.

Discussion

In this article we studied whether single neurons can solve linearly non-separable computational tasks, represented by the NFBP, by using a biophysically detailed multicompartment dSPN model. Based on the synaptic machinery of corticostriatal synapses onto dSPNs, we propose a learning rule that uses local synaptic calcium concentration and dopamine feedback signals: rewards for relevant stimuli and omitted rewards for irrelevant stimuli. Assuming first that single features in the NFBP are represented by clustered synapses, we show that the learning rule can solve the NFBP by strengthening (or stabilizing) synaptic clusters for relevant stimuli and weakening clusters for the irrelevant stimuli. The feature combinations for the relevant stimuli, stored in strengthened synaptic clusters, trigger supralinear dendritic responses in the form of plateau potentials, which is an important ingredient for the solution of the NFBP, as plateaus significantly increase the likelihood of neuronal spiking in SPNs in a robust way (Du et al., 2017).

The location of the synaptic clusters along the dendrites influenced the performance on the NFBP. In our model the region for optimal performance was 100–150 µm away from the soma, at about the same distance as where the somatic depolarization and induced spike probability, following activation of clustered synaptic input, was largest in the model by Lindroos and Hellgren Kotaleski (2021). Clusters placed further away produce smaller somatic depolarizations, due to dendritic filtering (Major et al., 2008), and as a consequence do not control the likelihood of somatic spiking as decisively. As the supralinear dendritic response is necessary for discriminating the relevant from the irrelevant stimuli, the performance with distally placed clusters decreases somewhat.

Earlier version: We further verified that supralinear dendritic responses were necessary to solve the NFBP by using randomly distributed synapses instead of clustered synapses for each feature. In this scenario, only one relevant stimulus is sometimes learned in a dendritic branch by the randomly distributed synapses. In the random setup the glutamate spillover models were updated to include the synapses in the whole dendritic branch, building on the notion that the single branch acts as a single computational unit (Losonczy and Magee, 2006; Branco and Häusser, 2010). The realism of this assumption about spillover is of course open to debate. In actual dendritic branches, however, the diffusion of signaling molecules likely plays a crucial role in both synaptic plasticity at existing synapses and for locally induced structural plasticity (Nishiyama and Yasuda, 2015; Chater et al., 2024). These processes may both enhance branch-specific supralinearities in qualitatively similar ways as when using spillover, but the mechanisms are not explored in the current model.

Revised version: We further verified that supralinear dendritic responses were necessary to solve the NFBP by using randomly distributed synapses instead of clustered synapses for each feature. In this scenario, only one relevant stimulus is sometimes learned in a dendritic branch by the randomly distributed synapses. In the random setup the glutamate spillover models were updated to include the synapses in the whole dendritic branch, building on the notion that the single branch acts as a single computational unit (Losonczy and Magee, 2006; Branco and Häusser, 2010). The realism of this spillover assumption remains open to debate. In actual dendritic branches, however, the diffusion of signaling molecules likely plays a crucial role in both synaptic plasticity at existing synapses and for locally induced structural plasticity (Nishiyama and Yasuda, 2015; Chater et al., 2024). These processes may both enhance branch-specific supralinearities in qualitatively similar ways as when using spillover, but the mechanisms are not explored in the current model.

Earlier version: By using a phenomenological inhibitory plasticity rule based on the BCM formalism (Bienenstock, Cooper and Munro, 1982), we also show that inhibitory synapses can significantly improve performance on the NFBP. This is in line with earlier theoretical studies where negative synaptic weights were required to solve the NFBP (Schiess, Urbanczik and Senn, 2016). In our setup with pre-existing synaptic clusters, inhibitory synapses made learning faster and increased performance by inhibiting supralinear NMDA responses for the irrelevant stimuli. This was specifically true in distal dendrites where the input impedance is higher (Branco, Clark and Häusser, 2010). The threshold for plateau potential initiation is also lower in distal dendrites compared to proximal (Losonczy and Magee, 2006) which likely will further extend the influence of inhibition in this region (Doron et al., 2017; Du et al., 2017). Similarly, in the scenario with distributed synapses, inhibition enables one of the relevant stimuli to be reliably encoded in the dendritic branch by strengthening the inhibitory synapses for features different from those of the encoded stimulus. Together, it therefore seems like inhibition not only has a role in learning (Chen et al., 2015; Cichon and Gan, 2015), but also improves the ability of the neuron to discriminate between stimuli with shared features.

Revised version: By using a phenomenological inhibitory plasticity rule based on the BCM formalism (Bienenstock et al., 1982), we also show that inhibitory synapses can significantly improve performance on the NFBP. This is in line with earlier theoretical studies where negative synaptic weights were required to solve the NFBP (Schiess et al., 2016). In our setup with pre-existing synaptic clusters, inhibitory synapses made learning faster and increased performance by inhibiting supralinear NMDA responses for the irrelevant stimuli. This was specifically true in distal dendrites where the input impedance is higher (Branco et al., 2010). The threshold for plateau potential initiation is also lower in distal dendrites compared to proximal (Losonczy and Magee, 2006) which likely will further extend the influence of inhibition in this region (Doron et al., 2017; Du et al., 2017). Similarly, in the scenario with distributed synapses, inhibition enables one of the relevant stimuli to be reliably encoded in the dendritic branch by strengthening the inhibitory synapses for features different from those of the encoded stimulus. Together, it therefore seems like inhibition not only has a role in learning (Chen et al., 2015; Cichon and Gan, 2015), but also improves the ability of the neuron to discriminate between stimuli with shared features.

Earlier version: In calcium-dependent plasticity models, a “no man’s land” has sometimes been proposed as a region where synaptic weights remain largely unchanged within a specific calcium range, reducing sensitivity to fluctuations and improving stability (Lisman, 2001; Moldwin, Azran and Segev, 2024). For our inhibitory rule, which builds on an BCM-like mechanism, introduction of such a mechanism in future work could make the system less sensitive to fluctuations in calcium concentration.

Revised version: In calcium-dependent plasticity models, a “no man’s land” has sometimes been proposed as a region where synaptic weights remain largely unchanged within a specific calcium range, reducing sensitivity to fluctuations and improving stability (Lisman, 2001; Moldwin et al., 2024). For our inhibitory rule, which builds on a BCM-like mechanism, introduction of such a mechanism in future work could make the system less sensitive to fluctuations in calcium concentration.

Earlier version: Although our learning rule infrequently solves the NFBP when used with only randomly distributed synapses, it can always learn to perform a linearly separable task, such as learning to respond to only one relevant stimulus (such as red strawberry). Moreover, the learning rule is general enough so that in addition to the feature-specific inputs related to the task, it can handle feature-unspecific inputs that might or might not be related to the NFBP. Finally, the learning rule is always “on”, continuously updating synapses with each stimulus presentation, which is a more realistic mechanism compared to using separate training and testing phases as in the field of machine learning. The synaptic weights automatically stabilize in the model when the performance improves, with metaplasticity playing a crucial role in this process. That is, rewards are then seen very regularly as the neuron has learned to spike for the relevant stimuli, while omitted rewards seldomly occur as the neuron stays silent when the irrelevant stimuli are provided.

Revised version: Although our learning rule infrequently solves the NFBP when used with only randomly distributed synapses, it can always learn to perform a linearly separable task, such as learning to respond to only one relevant stimulus (e.g.,red strawberry). Moreover, the learning rule is general enough so that in addition to the feature-specific inputs related to the task, it can handle feature-unspecific inputs that might or might not be related to the NFBP. Finally, the learning rule is always “on”, continuously updating synapses with each stimulus presentation, which is a more realistic mechanism compared to using separate training and testing phases as in the field of machine learning. The synaptic weights automatically stabilize in the model when the performance improves, with metaplasticity playing a crucial role in this process. That is, rewards are then seen very regularly as the neuron has learned to spike for the relevant stimuli, while omitted rewards rarely occur as the neuron stays silent when the irrelevant stimuli are provided.

When formulating the learning rule in this article, our goal was to base it on what is known regarding the synaptic machinery in corticostriatal synapses. This implied that the learning rule is based on the local calcium activity and on dopamine signals. Feedback from the dopamine system can also be viewed as an innate, evolutionarily encoded “supervisor”, which instructs neurons which feature combinations are beneficial and which ones should be avoided. However, in our case we do not use additional excitation to promote somatic spiking for only the relevant feature combinations, and in that sense the learning rule does not require a supervised learning paradigm. However, since the SPNs rest at very hyperpolarized membrane potential, our setup includes distributed excitatory inputs which are feature-unspecific in order to make sure that the neuron spikes for all stimuli, especially at the beginning of training. These additional inputs are on average weakened as learning progresses (as they are activated for all stimuli and thus often receive negative feedback). That general or noisy inputs are reduced during learning is in line with the observed reduction of execution-variability during motor-learning as a novice becomes an expert (Kawai et al., 2015). Specifically also, the underlying neuronal representation of corticostriatal synapses undergo a similar change during learning (Santos et al., 2015).

Since each synapse has its own calcium response, it is important for the learning rule to be able to follow individual synaptic activities. The LTP plasticity kernel in our model is for this reason itself plastic (metaplasticity), meaning that it changes its calcium dependence as a function of the history of reward and punishment. This setup helps the model separate clustered synaptic input from the feature-unspecific input at the beginning of training, as only clustered synapses will see enough calcium to fall within the LTP plastic range (LTP kernel).

We further use an asymmetric metaplasticity rule, where negative feedback causes a larger shift of the LTP kernel than a positive. This was necessary in order to prevent LTP in synapses that should ultimately undergo LTD. Hence, similarly to the classical loss-aversion tendency described in economic decision theory (Kahneman and Tversky, 1979), the model predicts that negative feedback will have a bigger impact in changing a well learned behavior on the single cell level than a positive feedback. Dopamine signaling has also been linked as a neural substrate to the decision making theory mentioned above (Stauffer et al., 2016).

Earlier version: It is not known whether single neurons solve the NFBP or other linearly non-separable tasks. However, many brain nuclei receive convergent inputs from numerous other brain nuclei, acting as integratory hubs (van den Heuvel and Sporns, 2013). Since feature binding evidently occurs in the brain, and functional clusters for single features such as visual stimulus orientation, receptive fields, color, or sound intensity exist on single neurons (Chen et al., 2011; Wilson et al., 2016; Iacaruso, Gasler and Hofer, 2017; Scholl, Wilson and Fitzpatrick, 2017; Ju et al., 2020), it is possible that the NFBP is a relevant task for neurons to solve. How brain regions with different synaptic machinery than the striatal dSPN might solve the NFBP remains a question, and reliance on other neuromodulatory signals may be part of the answer. For example, in the striatum, the indirect pathway SPNs (iSPN) have analogous synaptic machinery to the one in dSPNs, requiring calcium influx from the same sources for LTP and LTD, but are differently responsive to dopamine (Shen et al., 2008). In iSPNs a dopamine pause, together with a peak in adenosine, is required to trigger LTP, whereas a dopamine peak without peaks in adenosine rather promotes LTD (Shen et al., 2008; Nair et al., 2015). Therefore, we expect that an analogously formulated learning rule will also solve the NFBP in iSPNs, activating them for irrelevant feature combinations to e.g. suppress movement, and suppressing their activity for relevant feature combinations to facilitate movement. In addition, LTP in dSPNs is significantly facilitated by co-regulation of other neuromodulatory systems, such as when there is a coincident acetylcholine pause with the dopamine peak, which we have not explicitly included in the model (Nair et al., 2015; Bruce et al., 2019; Reynolds et al., 2022).

Revised version: It is not known whether single neurons solve the NFBP or other linearly non-separable tasks. However, many brain nuclei receive convergent inputs from numerous other brain nuclei, acting as integratory hubs (van den Heuvel and Sporns, 2013). Since feature binding evidently occurs in the brain, and functional clusters for single features such as visual stimulus orientation, receptive fields, color, or sound intensity exist on single neurons (Chen et al., 2011; Wilson et al., 2016; Iacaruso et al., 2017; Scholl et al., 2017; Ju et al., 2020), it is possible that the NFBP is a relevant task for neurons to solve. How brain regions with different synaptic machinery than the striatal dSPN might solve the NFBP remains a question, and reliance on other neuromodulatory signals may be part of the answer. For example, in the striatum, the indirect pathway SPNs (iSPN) have analogous synaptic machinery to the one in dSPNs, requiring calcium influx from the same sources for LTP and LTD, but are differently responsive to dopamine (Shen et al., 2008). In iSPNs a dopamine pause, together with a peak in adenosine, is required to trigger LTP, whereas a dopamine peak without peaks in adenosine rather promotes LTD (Shen et al., 2008; Nair et al., 2015). Therefore, we expect that an analogously formulated learning rule will also solve the NFBP in iSPNs, activating them for irrelevant feature combinations to e.g. suppress movement, and suppressing their activity for relevant feature combinations to facilitate movement. In addition, LTP in dSPNs is significantly facilitated by co-regulation of other neuromodulatory systems, such as when there is a coincident acetylcholine pause with the dopamine peak, which we have not explicitly included in the model (Nair et al., 2015; Bruce et al., 2019; Reynolds et al., 2022).

Added in the revised version: Methods and materials

Removed in the revised version: Methods

Earlier version: In this paper we introduce a local, calcium- and reward-based synaptic learning rule, constrained by experimental findings, to investigate learning in SPNs. The learning rule operates based on the changes in local calcium concentration resulting from synaptic activation patterns that evoke dendritic non-linearities, such as plateau potentials. Such events affecting the postsynaptic calcium concentration can enable learning of the nonlinear feature binding problem (NFBP). The learning rule is embedded in a biophysically detailed model of a dSPN built and simulated in the NEURON software, v8.2 (Carnevale et al., 2006). Here we will focus on the setup of the learning rule and only give a short summary of the neuron and synapse models, and emphasize changes compared to previously published versions. For a detailed description of the neuron model setup, see Lindroos et al., (2018) and Lindroos and Hellgren Kotaleski, (2021). The modeling of plateau potentials is explained in Trpevski et al., (2023). See also the section Code and Software Availability below.

Revised version: In this paper we introduce a local, calcium- and reward-based synaptic learning rule, constrained by experimental findings, to investigate learning in SPNs. The learning rule operates based on the changes in local calcium concentration resulting from synaptic activation patterns that evoke dendritic non-linearities, such as plateau potentials. Such events affecting the postsynaptic calcium concentration can enable learning of the nonlinear feature binding problem (NFBP). The learning rule is embedded in a biophysically detailed model of a dSPN built and simulated in the NEURON software, v8.2 (Carnevale and Hines, 2006). Here we will focus on the setup of the learning rule and only give a short summary of the neuron and synapse models, and emphasize changes compared to previously published versions. For a detailed description of the neuron model setup, see Lindroos et al. (2018); Lindroos and Hellgren Kotaleski (2021). The modeling of plateau potentials is explained in Trpevski et al. (2023). See also the section Code and Software Availability below.

Neuron model

In short, the dSPN model used here was taken from the library of biophysically detailed models in Lindroos and Hellgren Kotaleski (2021), including a reconstructed morphology and all of the most influential ion channels, including six calcium channels, each with its own voltage dependence and dendritic distribution. In accordance with Trpevski et al. (2023), the model was further extended with synaptic spines on selected dendrites. Each spine was modeled as two additional compartments consisting of a neck and a head region, and contains voltage-gated calcium channels of types R (Cav2.3), T (Cav3.2 and Cav3.3), and L (Cav1.2 and Cav1.3); the addition of explicit spines did not change the basic behavior of the model, such as the response to current injections, etc.

Calcium sources used in learning

Earlier version: The intracellular calcium concentration is separated into distinct pools that are used during the learning process. For learning in glutamatergic synapses, one pool for NMDA-evoked calcium concentration ([Ca]NMDA) is used, and another for L-type calcium concentration ([Ca]L-type), to reflect the different synaptic plasticity responses of the dSPN’s biochemical machinery to these two calcium sources in the corticostriatal synapse (Shen et al., 2008; Fino et al., 2010; Plotkin et al., 2013). Both pools are based on the calcium influx from the corresponding source (NMDA and the L-type channels Cav1.2 and Cav1.3, respectively). Inspired by findings that calcium from various voltage-dependent calcium channels affects GABAergic plasticity in different systems (Kurotani et al., 2008; Hulme and Connelly, 2014; Udakis et al., 2020), we also incorporate a third calcium pool based on calcium influx from all voltage gated channels (T-, R-, L- and N-type, [Ca]V). All pools include extrusion mechanisms in the form of a calcium pump as well as a one-dimensional time-decay. The calcium pump follows the implementation in Wolf et al., (2005). The parameters for the [Ca]NMDA model were manually tuned to match the [Ca] amplitudes and durations reported in Dorman et al., (2018). The voltage gated calcium channel conductances in the spines were also manually tuned to match the relative calcium proportions in Carter and Sabatini, (2004) and Higley and Sabatini, (2010), as well as the calcium amplitudes due to stimulation with backpropagating action potentials (Shindou, Ochi-Shindou and Wickens, 2011). Spatial calcium diffusion was not included in the present model.

Revised version: The intracellular calcium concentration is separated into distinct pools that are used during the learning process. For learning in glutamatergic synapses, one pool for NMDA-evoked calcium concentration ([Ca]NMDA) is used, and another for L-type calcium concentration ([Ca]L−type), to reflect the different synaptic plasticity responses of the dSPN’s biochemical machinery to these two calcium sources in the corticostriatal synapse (Shen et al., 2008; Fino et al., 2010; Plotkin et al., 2013). Both pools are based on the calcium influx from the corresponding source (NMDA and the L-type channels Cav1.2 and Cav1.3, respectively). Inspired by findings that calcium from various voltagedependent calcium channels affects GABAergic plasticity in different systems (Kurotani et al., 2008; Hulme and Connelly, 2014; Udakis et al., 2020), we also incorporate a third calcium pool based on calcium influx from all voltage gated channels (T-, R-, L- and N-type, [Ca]V). All pools include extrusion mechanisms in the form of a calcium pump as well as a one-dimensional time-decay. The calcium pump follows the implementation in Wolf et al. (2005). The parameters for the [Ca]NMDA model were manually tuned to match the [Ca] amplitudes and durations reported in Dorman et al.(2018). The voltage gated calcium channel conductances in the spines were also manually tuned to match the relative calcium proportions in Carter and Sabatini (2004)and Higley and Sabatini (2010), as well as the calcium amplitudes due to stimulation with backpropagating action potentials (Shindou et al., 2011). Spatial calcium diffusion was not included in the present model.

Glutamatergic synaptic input

Clustered synapses were modeled following Trpevski et al. (2023) using a version of the saturating synapse model in Destexhe et al. (1994). The synaptic model included AMPA and NMDA conductances activated on spines as well as extrasynaptic NMDA conductances located on the dendritic shafts adjacent to the spine. The strength of the individual synapses on spines were scaled using a weight parameter (w) taking on positive values. More explicitly, in our model, the parameter w represents the percentage of the maximum synaptic conductance gmax. For instance, when w = 0.5, the synaptic conductance (g) is 0.5 ⋅ gmax. Importantly, changes in w affect both NMDA and AMPA receptor conductances simultaneously, as both share the same scaling factor. Extrasynaptic conductances were activated following sufficiently large stimulation of nearby spines, resulting in glutamate spillover. Extrasynaptic NMDA synapses were included because they provide robust all-or-none dendritic plateaus (Trpevski et al., 2023). The maximal conductance of both synaptic and extrasynaptic NMDARs was set to 2.5 nS, which is lower than the 3.5 nS used in Trpevski et al. (2023). This reduction in conductance was implemented to compensate for the increased number of inputs in the model. In this model, spillover was triggered when the sum of the weights of active synapses in a cluster exceeded a threshold value of 2. For instance, eight synapses with individual weights greater than 0.25 would activate spillover. This threshold was lower than the value of 4 used in Trpevski et al. (2023) in order to facilitate spillover and enhance non-linearity in the context of distributed synapses. However, this scaling only affects the size of the clustered inputs presumed to give rise to plateau potentials, which likely varies across different neuronal subtypes or even between distinct dendritic branches within the same subtype. In the NEURON simulation environment, an integrate-and-fire cell was used to generate a spike in extrasynaptic NMDARs when the spillover threshold is reached. As synaptic weight increases, fewer synapses are required to meet the threshold. This also links long-term potentiation (LTP) to enhanced glutamate spillover. In addition to task-related inputs, the neuron also receives background noise, modeled as excitatory and inhibitory synapses distributed across the dendrites. Since the number of electrical compartments in the model is significantly lower than the number of synapses in a real SPN, these converging synapses were represented by a single non-saturating synapse with input frequency scaled accordingly. This was implemented using a dual exponential synapse model based on the NEURON exp2syn mechanism following Lindroos and Hellgren Kotaleski (2021) with time constants and other parameters adapted from Hjorth et al. (2020).

Learning rule

Excitatory synaptic plasticity

The learning rule is grounded in experimental evidence indicating that striatal LTP relies on NMDA receptor activation and the presence of dopamine, whereas LTD depends on the activation of L-type calcium channels (Cav1.3) and mGluR5 receptors in the absence of dopamine (Shen et al., 2008; Fino et al., 2010; Plotkin et al., 2013; Yagishita et al., 2014; Fisher et al., 2017; Shindou et al., 2019). Under basal dopamine levels, no significant synaptic plasticity is assumed to occur. This rule outlines a reward-based learning mechanism where a dopamine peak enhances synaptic weight via an LTP plasticity kernel, while a dopamine pause reduces synaptic weight through an LTD kernel. Thus, dopamine functions as a switch, determining whether LTP or LTD pathways are activated (as illustrated in Figure 1D). Notably, synaptic weights affect both AMPA and NMDA receptors, consistent with findings that the NMDA-to-AMPA ratio remains unchanged following LTP (Watt et al., 2004). Additionally, the previously described spillover mechanism is linked to synaptic weights, representing an attempt to associate plasticity with increased glutamate spillover resulting from the retraction of astrocytic processes from the spine (Henneberger et al., 2020). The excitatory synaptic plasticity rule encompasses key components such as LTP, LTD, and metaplasticity, each of which will be elaborated upon in the subsequent sections.

LTP

Earlier version: The LTP process is initiated by elevated dopamine levels, following calcium influx through NMDA receptors. Synaptic strength is adjusted using the LTP kernel, a bell-shaped function where the maximum increase occurs when peak NMDA calcium levels are at the kernel’s midpoint; deviations above or below this midpoint result in smaller increases (Figure 7A). The left side of the LTP kernel represents the lower calcium threshold necessary for LTP induction, while the right side indicates an upper calcium threshold beyond which LTP does not occur. This framework, in conjunction with metaplasticity, establishes an upper boundary for synaptic strength, ensuring that synaptic weights do not exceed a certain limit (Zenke and Gerstner, 2017; Zenke, Gerstner and Ganguli, 2017). Since each synapse experiences unique calcium levels (due to different synaptic conductance, local input resistance, etc), adjusting the calcium dependence of the LTP kernel allows for individualized tuning of synaptic weights (see Figure 7A and 7C for illustrations).

Revised version: The LTP process is initiated by elevated dopamine levels, following calcium influx through NMDA receptors. Synaptic strength is adjusted using the LTP kernel, a bell-shaped function where the maximum increase occurs when peak NMDA calcium levels are at the kernel’s midpoint; deviations above or below this midpoint result in smaller increases (Figure 7A, top panel). The left side of the LTP kernel represents the lower calcium threshold necessary for LTP induction, while the right side indicates an upper calcium threshold beyond which LTP does not occur. This framework, in conjunction with metaplasticity, establishes an upper boundary for synaptic strength, ensuring that synaptic weights do not exceed a certain limit (Zenke et al., 2017; Zenke and Gerstner, 2017).

Synaptic Plasticity Rules: Calcium and Dopamine Interactions in Synaptic Weight Modification

Removed in the revised version: (A) : Synaptic weight updates following a dopamine peak. (Left) LTP kernel is a bell-shaped curve describing the amount of weight increase, which happens over a region of [Ca]NMDA. (Right) A wider bell-shaped kernel, i.e. the metaplasticity kernel, determines how the LTP kernel is shifted along the calcium level ([Ca]NMDA) axis following a peak in dopamine.

Removed in the revised version: (B) : Synaptic weight updates following a dopamine pause. (Left) The LTD plasticity kernel. The LTD threshold is constant and set at 70 nM (Right). Metaplasticity describing how the LTP kernel shifts along the calcium axis following a dopamine pause.

Removed in the revised version: (C) : A schematic of how the LTP kernel is shifted following a dopamine peak, and examples of NMDA calcium levels before and after the shift. The NMDA calcium levels increase following activation of the strengthened synapse due to the previous dopamine peak (illustrated with the red circle jumping to the blue circle). The LTP kernel also moves downwards, and in this example the strengthened synapse (blue circle) would stabilize, and thus would not increase in strength even though additional rewards arrive.

Removed in the revised version: (D) : A schematic showing how the LTP kernel is shifted following a dopamine pause, and examples of NMDA calcium levels before and after the shift. The calcium level shortly before the dopamine pause is shown with the red circle. The dopamine pause causes a decrease in synaptic weight and a smaller calcium response results in the next synapse activation (illustrated with the red circle jumping to the blue one). Also, the LTP kernel moves up towards higher NMDA calcium levels, and thus the weakened synapse (blue circle) is more likely to stay out of the plastic region of the LTP kernel (unless it in the future is activated with dopamine peaks more regularly than with dopamine pauses).

Earlier version: (E) : Illustration of the inhibitory plasticity rule. (Left) Changes in synaptic weight for active (beige) and inactive (blue) synapses based on voltage dependent calcium levels in the dendritic shaft at the location of the inhibitory synapse. The dashed lines show the minimum/lower (CTL) and maximum/higher (CTH) threshold. (Right) Functions for updating the lower and higher thresholds with voltage dependent calcium level. The asterisk denotes the Ca level where the curves for active and inactive synapses meet, which is a point of zero update.

Revised version: (A) : Synaptic weight updates following a dopamine peak. (Top) The LTP kernel is a bell-shaped curve describing the amount of weight increase, which happens over a region of [Ca]NMDA. (Bottom) A wider bell-shaped kernel, i.e., the metaplasticity kernel, determines how the LTP kernel is shifted along the calcium level ([Ca]NMDA) axis following a peak in dopamine. (B) : Synaptic weight updates following a dopamine pause. (Top) The LTD plasticity kernel. The LTD threshold is constant and set at 70 nM. (Bottom) Metaplasticity describing how the LTP kernel shifts along the calcium axis following a dopamine pause. (C) : A schematic of how the LTP kernel is shifted following a dopamine peak (left) and pause (right), together with examples of NMDA calcium levels before and after the shift. The NMDA calcium levels change following the potentiation or suppression of the synapse (illustrated with the red circle jumping to the blue circle). The LTP kernel also moves as indicated by the arrows and colors following a peak or a pause, respectively. (D) : Illustration of the inhibitory plasticity rule. (Top) Changes in synaptic weight for active (beige) and inactive (blue) synapses based on voltage-dependent calcium levels in the dendritic shaft at the location of the inhibitory synapse. The dashed lines show the minimum/lower (θinh,low) and maximum/higher (θinh,high) thresholds. (Bottom) Functions for updating the thresholds of the upper panel, depending on the voltage-dependent calcium level. The asterisk denotes the calcium level where the curves for active and inactive synapses meet, which is a point of zero update.

Added in the revised version: Since each synapse experiences unique calcium levels (due to different synaptic conductance, local input resistance, etc), adjusting the calcium dependence of the LTP kernel allows for individualized tuning of synaptic weights (see Figure 7A and C for illustrations).

Earlier version: The update of the synaptic weights follow a rule formulated using the derivative of the sigmoid function. The sigmoid function (Eq. 1) transforms the calcium concentration into a normalized range between 0 and 1: where θ represents the midpoint of the transition, representing the calcium concentration at which the function increases most rapidly, and β determines the steepness of this transition.

Revised version: The update of the synaptic weights follow a rule formulated using the derivative of the sigmoid function (Equation 2). The sigmoid function (Equation 1) transforms the calcium concentration into a normalized range between 0 and 1: where θ represents the midpoint of the transition, representing the calcium concentration at which the function increases most rapidly, and β determines the steepness of this transition.

The derivative of the sigmoid function, given by: produces a bell-shaped curve, which forms the LTP kernel. This kernel peaks at the midpoint θ, where synaptic modification is maximized, and decays symmetrically on each side. The parameterβ here controls the width of the kernel, thereby regulating the sensitivity of synaptic weight changes to calcium fluctuations.

Earlier version: Based on this general description of the mathematical formulas underlying the LTP kernel, the LTP learning rule for the synaptic weights is formulated as: where dwltp represents the change in synaptic weight, ηltp is the learning rate, and [Ca]NMDA denotes the calcium concentration mediated by NMDA receptors. The parameters θltp and βltp define the midpoint and slope of the sigmoid curve specific to LTP. The parameter values used in these equations are provided in Table 1.

Revised version: Based on this general description of the mathematical formulas underlying the LTP kernel, the LTP learning rule for the synaptic weights is formulated as: where Δwltp represents the change in synaptic weight, ηltp is the learning rate, and [Ca]NMDA denotes the calcium concentration mediated by NMDA receptors. The parameters θltp and βltp define the midpoint and slope of the sigmoid curve specific to LTP. The parameter values used in these equations are provided in Table 1).

LTD

Earlier version: The LTD process is triggered by a dopamine pause and is dependent on L-type calcium. The LTD plasticity kernel (Eq. 4), describes a threshold level of the calcium level necessary for LTD to occur (Shindou, Ochi-Shindou and Wickens, 2011), after which the decrease in synaptic weight is linearly proportional to the amplitude of the peak calcium level, and is scaled by the learning rate (ηltd, see left panel of Figure 7B for an illustration). The peak calcium threshold is implemented with a sigmoid function, whose slope parameter, βltd, was set to a high-value to make the curve resemble a step function. The threshold is set at 70 nM and is fixed. The NMDA calcium dependence of the LTP kernel midpoint, θltp, is also increased during LTD, as described in the next section (Figure 7B right panel and Figure 7D).

Revised version: The LTD process is triggered by a dopamine pause and is dependent on L-type calcium. The LTD plasticity kernel (Equation 4; see Figure 7B, top panel) describes a threshold calcium level necessary for LTD to occur (Shindou et al., 2011). Once this threshold is exceeded, the decrease in synaptic weight is linearly proportional to the amplitude of the peak calcium level and is scaled by the learning rate (ηltd). The peak-calcium threshold is implemented with a sigmoid function whose slope parameter (βltd) is set high so the curve approximates a step function. The threshold is fixed at 70 nM. The NMDA-calcium dependence of the LTP kernel is also increased during LTD, as described in the next section (Figure 7B, bottom panel, and Figure 7C, right panel).

Metaplasticity

Earlier version: Metaplasticity is a form of regulatory mechanism changing the state of the synapses in such a way as to influence subsequent learning (Abraham, 2008). In our model we implemented this as a reward-dependent change of the calcium concentration over which the LTP kernel was operating. Specifically, activation of the LTD pathway triggered a shift of the LTP kernel towards higher calcium concentrations while activation of the LTP pathway pushed it in the opposite direction. Metaplasticity was also implemented using a bell-shaped kernel with the same midpoint as the LTP kernel, but with a wider calcium dependence (Figures 7A-D). The position of the metaplasticity kernel is updated at the same time with the LTP kernel, so that both kernels remain centered at the same calcium level (Figure. 7C and 7D). This setup with the two dynamically moving kernels together allowed for a wide range of calcium levels to induce plasticity. In addition, using a wider metaplasticity kernel that:

Revised version: Metaplasticity is a form of regulatory mechanism changing the state of the synapses in such a way as to influence subsequent learning (Abraham, 2008). In our model we implemented this as a reward-dependent change of the calcium concentration over which the LTP kernel was operating. Specifically, activation of the LTD pathway triggered a shift of the LTP kernel towards higher calcium concentrations while activation of the LTP pathway pushed it in the opposite direction.

Added in the revised version: Metaplasticity was also implemented using a bell-shaped kernel with the same midpoint as the LTP kernel, but with a wider calcium dependence. The position of the metaplasticity kernel is updated at the same time with the LTP kernel, so that both kernels remain centered at the same calcium level (Figure 7A–C). This setup with the two dynamically moving kernels together allowed for a wide range of calcium levels to induce plasticity.

Added in the revised version: In addition, using a wider metaplasticity kernel ensured the following:

Earlier version: Following activation of the LTP pathway (i.e. after a dopamine peak, see Figure 7C for an illustration),

Revised version: Following activation of the LTP pathway (i.e., after a dopamine peak, see Figure 7C, left panel for an illustration):

In synapses with low calcium levels, the kernel will be shifted closer to the observed calcium level and thereby eventually enable LTP in synapses that are regularly activated during rewards, despite that they initially don’t generate big elevations in calcium.

Earlier version: In synapses with already high calcium levels (above the kernel midpoint), the kernel will be shifted away from the observed calcium level, and thereby protect the synapse from excessive LTP and instead stabilize the weight (as illustrated with the red circle moving to the blue one in Figure 7C).

Revised version: In synapses with already high calcium levels (above the kernel midpoint), the kernel will be shifted away from the observed calcium level, and thereby protect the synapse from excessive LTP and instead stabilize the weight (as illustrated with the red circle moving to the blue one in Figure 7C,left panel).

Earlier version: Following activation of the LTD pathway (i.e. after a dopamine pause), the kernel will be shifted towards higher calcium levels and thereby reduce the likelihood of LTP in these synapses that are weakened following LTD (see Figure 7D for an illustration).

Revised version: Following activation of the LTD pathway (i.e., after a dopamine pause), the kernel will be shifted towards higher calcium levels and thereby reduce the likelihood of LTP in these synapses that are weakened following LTD (see Figure 7C, right panel for an illustration).

Added in the revised version: The update of the LTP kernel was further asymmetric following activation of the LTD or LTP pathways in such a way that the LTD pathway caused a larger shift of the metaplasticity kernel than the LTP pathway. This further reduced the likelihood of inducing LTP in synapses often participating in LTD or in synapses randomly activated with regard to the dopamine feedback signal.

Added in the revised version: It also allows initially weaker synapses to be recruited for LTP if they are more consistently coactive with a reward than with an omitted reward (an illustration of this can be seen in Figure 5E, where one of the yellow synapses is strengthened in the middle of the training session).

Earlier version: The update of the LTP kernel was further asymmetric following activation of the LTD or LTP pathways in such a way that the LTD pathway caused a larger shift of the metaplasticity kernel than the LTP pathway. This further reduced the likelihood of inducing LTP in synapses often participating in LTD or in synapses randomly activated with regard to the dopamine feedback signal. It also allows initially weaker synapses to be recruited for LTP if they are more consistently co-active with a reward than with an omitted reward (an illustration of this can be seen in Figure 5E, where one of the yellow synapses is strengthened in the middle of the training session). The LTP kernel was updated according to Eq. 5, based on the learning rate (ηs), the peak NMDA calcium level ([Ca]NMDA), and the slope of the sigmoid curve (βmp). The value of the slope βmp creates a wider bell-shaped function in Eq. 5. The rate ηs, captures the described asymmetry in how the metaplasticity kernel is shifted in response to dopamine feedback where the specific rates following dopamine peaks and pauses are denoted ns,ltp and ns,ltp, respectively. Here ns,ltd is four times larger than ns,ltp (see Table 1). For better clarity, we also provide the pseudo code for the learning rule below: Inputs:

Revised version: The LTP kernel was updated according to Equation 5, based on the learning rate (ηs), the peak NMDA calcium level ([Ca]NMDA), and the slope of the sigmoid curve (βmp). The value of the slope βmp creates a wider bell-shaped function in Equation 5. The rate ηs captures the described asymmetry in how the metaplasticity kernel is shifted in response to dopamine feedback, where the specific rates following dopamine peaks and pauses are denoted ηs,ltp and ηs,ltd, respectively. Here, ηs,ltd is four times larger than ηs,ltp (see Table 1).

Added in the revised version: Pseudocode for the Excitatory Plasticity Rule

Added in the revised version: Inputs

[Ca]NMDA : Calcium level associated with the NMDA channel

[Ca]Ltype : Calcium level associated with the L-type channel

dopamine_level: Dopamine level (1 for peak, -1 for pause, 0 for basal)

Output

Δw: Change in synaptic weight

Added in the revised version: Algorithm

• begin

‐if dopamine_level == 1 then

Removed in the revised version: # Calculate the change in synaptic weight for the LTP process

Removed in the revised version: Δw = solve Eq. 3 with [Ca]NMDA, ηltp, θltp, and βltp

Earlier version: # Update the LTP kernel midpoint for metaplasticity update θltp

Revised version: ‐# Compute LTP weight change and adjust kernel midpoint for metaplasticity

Added in the revised version: *Δw = solve Equation 3with [Ca]NMDA, ηltp, θltp, and βltp

Earlier version: using Eq. 5 with [Ca]NMDA, ηs,ltp, and βmp

Revised version: *update θltp using Equation 5with [Ca]NMDA, ηs,ltp, βmp

‐else if dopamine_level == -1 then

Added in the revised version: ‐# Compute LTD weight change and adjust LTP kernel midpoint for metaplasticity

Removed in the revised version: # Calculate the change in synaptic weight for the LTD process

Earlier version: Δw = solve Eq. 4 with [Ca]Ltype, ηltd, θltd, and βltd

Revised version: *Δw = solve Equation 4with [Ca]Ltype, ηltd, θltd, and βltd

Removed in the revised version: # Update the LTP kernel midpoint for metaplasticity

Earlier version: update θltp using Eq. 5 with [Ca]NMDA, ηs,ltd and βmp

Revised version: *update θltp using Equation 5with [Ca]NMDA, ηs,ltd, βmp

‐else if dopamine_level == 0 then

Earlier version: # No significant synaptic plasticity under basal dopamine levels

Revised version: ‐# No plasticity under basal dopamine

* Δw = 0

Removed in the revised version: end if

‐return Δw

• end

Inhibitory Synaptic Plasticity

Earlier version: In contrast to the well studied mechanistic underpinnings of glutamatergic plasticity in e.g. SPNs, much less is known about how inhibitory synapses might be updated during learning. The inhibitory plasticity rule developed here is therefore more phenomenological and exploratory in nature and was developed to enhance nonlinearities in the local dendrite. The rule is based on the Bienenstock-Cooper-Munro (BCM) formalism (Bienenstock, Cooper and Munro, 1982). In the BCM rule there is a threshold level of synaptic activity below which LTD is triggered and above which LTP is triggered (Figure 7E). Inspired by (Gandolfi et al., 2020; Chapman, Nuwer and Jacob, 2022; Ravasenga et al., 2022), we use calcium from all voltage-gated calcium channels ([Ca]V) as the indicator of excitatory synaptic activity near the dendritic shaft where an inhibitory synapse is located. The inhibitory rule is designed to passively observe and respond to the surrounding excitatory synaptic activity, governed by local voltage-gated calcium influx, but without reliance on dopamine or explicit feedback. It operates at a slower pace to ensure that it follows the average excitatory activity levels, and is meant to enhance the contrast in activity by amplifying local differences in excitatory synaptic efficacy.

Revised version: In contrast to the well studied mechanistic underpinnings of glutamatergic plasticity in e.g. SPNs, much less is known about how inhibitory synapses might be updated during learning. The inhibitory plasticity rule developed here is therefore more phenomenological and exploratory in nature and was developed to enhance nonlinearities in the local dendrite. The rule is based on the Bienenstock-Cooper-Munro (BCM) formalism (Bienenstock et al., 1982). In the BCM rule there is a threshold level of synaptic activity below which LTD is triggered and above which LTP is triggered (Figure 7D). Inspired by Gandolfi et al. (2020); Chapman et al. (2022); Ravasenga et al. (2022), we use calcium from all voltage-gated calcium channels ([Ca]V) as the indicator of excitatory synaptic activity near the dendritic shaft where an inhibitory synapse is located. The inhibitory rule is designed to passively observe and respond to the surrounding excitatory synaptic activity, governed by local voltage-gated calcium influx, but without reliance on dopamine or explicit feedback. It operates at a slower pace to ensure that it follows the average excitatory activity levels, and is meant to enhance the contrast in activity by amplifying local differences in excitatory synaptic efficacy.

Added in the revised version: Equation 6describes how the inhibitory synaptic weights change based on the calcium concentration (calculated from calcium influx from all voltage-gated ion channels close to the synapse).

Earlier version: Equation 6 describes how the inhibitory synaptic weights change based on the calcium concentration (calculated from calcium influx from all voltage-gated ion channels close to the synapse). It consists of two sigmoidal terms, where the parameters a and b determine the contribution of each component, βinh determines the steepness of the transitions and CTL and CTH determines the calcium concentration at which each sigmoid is half activated (see Figure 7E). Based on this, and modulated by a rate constant, ηact, Equation 7 describes the specific rule used to update the weights of the inhibitory synapses. The weight change also depends on the magnitude of the weights themselves, in such a way that intermediate weights update faster than small or large weights. This helps stabilize the weights of the inhibitory synapses at the end of learning, and prevents the weight from taking on large or negative values. Large weights are then close to at the end of learning and small weights close to 0 (see Figure 5D for an example and table 2 for parameter values).

Revised version: It consists of two sigmoidal terms, where the parameters a and b determine the contribution of each component, βinh determines the steepness of the transitions and θinh,low and θinh,high determines the calcium concentration at which each sigmoid is half activated (see Figure 7D). Based on this, and modulated by a rate constant, ηact, Equation 7describes the specific rule used to update the weights of the inhibitory synapses.

Added in the revised version: The weight change also depends on the magnitude of the weights themselves, in such a way that intermediate weights update faster than small or large weights. This helps stabilize the weights of the inhibitory synapses at the end of learning, and prevents the weight from taking on large or negative values. Large weights are then close to at the end of learning and small weights close to 0 (see Figure 5D for an example and Table 2for parameter values).

Earlier version: The learning rule further differentiates between active and inactive synapses. For active synapses: i) if the calcium concentration exceeds CTH, the weight is depressed, ii) if calcium falls between CTL and CTH, the weight is potentiated and iii) when calcium is below CTL, little or no weight change occurs. For inactive synapses, the rule is reversed: i) when calcium exceeds CTH, the weight is potentiated, ii) if calcium falls between CTL and CTH, the weight is depressed and iii) no changes occur when calcium is below CTL (as in the case of active synapses). The schematic in Figure 7E (left panel) illustrates these weight updates, showing how the learning rule differentially affects active and inactive synapses. The black curves represent the relationship between calcium concentration and synaptic weight change, with the beige and blue shaded regions indicating potentiation and depression zones for each case.

Revised version: The learning rule further differentiates between active and inactive synapses. For active synapses: i) if the calcium concentration exceeds θinh,high, the weight is depressed, ii) if calcium falls between θinh,low and θinh,high, the weight is potentiated and iii) when calcium is below θinh,low, little or no weight change occurs. For inactive synapses, the rule is reversed: i) when calcium exceeds θinh,high, the weight is potentiated, ii) if calcium falls between θinh,low and θinh,high, the weight is depressed and iii) no changes occur when calcium is below θinh,low (as in the case of active synapses). The schematic in Figure 7D (top panel) illustrates these weight updates, showing how the learning rule differentially affects active and inactive synapses. The black curves represent the relationship between calcium concentration and synaptic weight change, with the beige and blue shaded regions indicating potentiation and depression zones for each case.

Earlier version: The half activation thresholds, CTL, and CTH, are further dynamically updated to track the calcium levels of the synapse. These updates of the thresholds ensure that the plasticity rule remains sensitive to changes in calcium concentration and stabilize synaptic learning. Equations 8 and 9 describe these adaptive dynamics. This is, in fact, metaplasticity of the inhibitory synapses.

Revised version: The half activation thresholds, θinh,low, and θinh,high, are further dynamically updated to track the calcium levels of the synapse. These updates of the thresholds ensure that the plasticity rule remains sensitive to changes in calcium concentration and stabilize synaptic learning.Equation 8 and Equation 9describe these adaptive dynamics. This is, in fact, metaplasticity of the inhibitory synapses.

Earlier version: The upper threshold CTH shifts towards the highest observed calcium level in the synapse, while the lower threshold CTL moves to a level below the maximum. The rate of these adjustments is determined by the parameters ηTH and ηTL, which scale the influence of calcium activity on the threshold dynamics. The parameters aTH, bTH determine the rate of change of CTH, while aTL, bTL define the rate of change of CTL (see Table 2 for parameter values). The parameter c in Equation 9 is used to make the two curves for the high and low thresholds (CTL, CTH) to intersect at zero which minimizes the fluctuations around the transition point (green asterisk in Figure 7E, right panel).

Revised version: The upper threshold θinh,high shifts towards the highest observed calcium level in the synapse, while the lower threshold θinh,low moves to a level below the maximum. The rate of these adjustments is determined by the parameters ηinh,high and ηinh,low, which scale the influence of calcium activity on the threshold dynamics. The parameters ainh,high, binh,high determine the rate of change of θinh,high, while ainh,low, binh,low define the rate of change of θinh,low (see Table 2for parameter values). The parameter c in Equation 9is used to make the two curves for the high and low thresholds (θinh,low, Cinh,high) intersect at zero which minimizes the fluctuations around the transition point (green asterisk in Figure 7D, bottom panel).

Earlier version: Pseudocode for the algorithm describing inhibitory plasticity is given below.

Revised version: Pseudocode for the Inhibitory Plasticity Rule

Inputs

[Ca]V : Ca concentration from voltage-gated calcium channels

Output

Earlier version: - dwinh: Change in synaptic weight

Revised version: return Δwinh: Change in synaptic weight

Added in the revised version: Algorithm

•begin

‐ if synapse_active then

Earlier version: # Active synapse: Compute weight change using negative learning rate

Revised version: ‐ # Active synapse: Compute weight change using negative learning rate (ηact) inh

Added in the revised version: *Δwinh = solve Equation 7 with [Ca]V, ηact, θinh,high, θinh,low, ainh, binh,

Added in the revised version: *update θinh,high using Equation 8 with θinh,high, θinh,low, ηinh,high, ainh,high,binh,high

Added in the revised version: *update θinh,low using Equation 9 with θinh,high, θinh,low, ηinh,low,ainh,low, binh,low, c

Removed in the revised version: dw = solve Eq.7 with [Ca], ηact, CTH, CTL, ηTH, ainh, binh,

Removed in the revised version: # Update the upper threshold (CTH) based on Eq. 8

Removed in the revised version: update CTH using Eq. 8 with CTH, CTL, ηTH, aTH,bTH

Removed in the revised version: # Update the lower threshold (CTL) based on Eq. 9

Removed in the revised version: update CTL using Eq. 9 with CTH, CTL, ηTL, aTL,bTL, c

‐else

Earlier version: # Inactive synapse: Compute weight change using positive learning rate

Revised version: ‐# Inactive synapse: Compute weight change using positive learning rate (ηact)

Added in the revised version: *Δwinh = solve Equation 7with [Ca]V, ηact, θinh,high, θinh,low, ainh, binh,

Added in the revised version: *update θinh,high using Equation 8with θinh,high, θinh,low, ηinh,high, ainh,high, binh,high

Added in the revised version: *update θinh,low using Equation 9with θinh,high, θinh,low, ηinh,low, ainh,low, binh,low, c

Added in the revised version: ‐return Δwinh

Removed in the revised version: dw = solve Eq.7 with [Ca]V, ηact, CTH, CTL, ηTH, ainh, binh,

Removed in the revised version: # Update the upper threshold (CTH) based on Eq. 8

Removed in the revised version: update CTH using Eq. 8 with CTH, CTL, ηTH, aTH,bTH

Removed in the revised version: # Update the lower threshold (CTL) based on Eq. 9

Removed in the revised version: update CTL using Eq. 9 with CTH, CTL, ηTL, aTL,bTL, c

Removed in the revised version: end if

Removed in the revised version: return dwinh

•end

Training Procedure

Earlier version: The features of a stimulus were represented by the shape and color of bananas and strawberries. We presented the neuron with a sequence of stimuli, typically around 960, each belonging to one of four possible feature combinations, as illustrated in Figure 1. The four stimuli: ‘red strawberry’, ‘yellow banana’, ‘red banana’ and ‘yellow strawberry’, were presented in random order three times each within a block of 12 stimuli, followed by another reshuffled block of 12, and so on. Each stimulus presentation lasted 20 ms, during which time all the stimuli-related synapses receive one randomly-timed spike per synapse. Whenever relevant, inhibitory synapses were activated concurrently with the excitatory synapses and the stimuli lasted for 100 ms. The stimuli presentation was followed by a reward cue lasting 50 ms that arrived 300 ms after the stimulus onset. The reward cue was represented with +1 for relevant stimuli if the neuron spiked (representing a peak in dopamine), with −1 for irrelevant stimuli if the neuron spiked (representing a dopamine pause) and 0 for baseline levels or when the neuron was silent. However, in an additional subthreshold learning task, the reward cue was delivered even in the absence of somatic spiking for active glutamate synapses, such that if they were active following the relevant stimuli (‘red strawberry’ and ‘yellow banana’) they got a dopamine peak, and a dopamine pause for irrelevant stimuli (Figure 6—figure supplement 1B and 1C). The time between two stimuli was 800 ms, long enough to allow for the voltage, calcium (and all other state variables in the model) to return to their baseline values (Figure 3A). During the stimulus presentation, the stimuli-related synapses receive one randomly-timed spike per synapse. Further, the learning rule was on during the whole procedure, such that synapses were continuously updated throughout the simulation. Hence, there were no separate training and testing phases in which synapses were plastic and frozen, respectively.

Revised version: The features of a stimulus were represented by the shape and color of bananas and strawberries. We presented the neuron with a sequence of stimuli, typically around 960, each belonging to one of four possible feature combinations, as illustrated in Figure 1. The four stimuli: ‘red strawberry’, ‘yellow banana’, ‘red banana’ and ‘yellow strawberry’, were presented in random order three times each within a block of 12 stimuli, followed by another reshu?ed block of 12, and so on. Each stimulus presentation lasted 20 ms, during which time all the stimuli-related synapses receive one randomly-timed spike per synapse(see Figure 3B). Whenever relevant, inhibitory synapses were activated concurrently with the excitatory synapses and the stimuli lasted for 100 ms. The stimuli presentation was followed by a reward cue lasting 50 ms that arrived 300 ms after the stimulus on-set. The reward cue was represented with +1 for relevant stimuli if the neuron spiked (representing a peak in dopamine), with -1 for irrelevant stimuli if the neuron spiked (representing a dopamine pause) and 0 for baseline levels or when the neuron was silent. However, in an additional subthreshold learning task, the reward cue was delivered even in the absence of somatic spiking for active glutamate synapses, such that if they were active following the relevant stimuli (‘red strawberry’ and ‘yellow banana’) they got a dopamine peak, and a dopamine pause for irrelevant stimuli (Figure 6—figure Supplement 1B and C). The time between two stimuli was 800 ms, long enough to allow for the voltage, calcium (and all other state variables in the model) to return to their baseline values (Figure 3B). During the stimulus presentation, the stimuli-related synapses receive one randomly-timed spike per synapse. Further, the learning rule was on during the whole procedure, such that synapses were continuously updated throughout the simulation. Hence, there were no separate training and testing phases in which synapses were plastic and frozen, respectively.

Removed in the revised version: Two distinct synaptic arrangements were used for representing the features: the clustered and the distributed setups.

Clustered setup

Earlier version: This setup was based on the assumption of pre-existing clustered synapses for each feature. Features were allocated to two dendritic branches. Each branch had each feature represented with five synapses clustered closely on a single dendrite. Depending on the feature combination, two, three, or four features were represented in clusters on one or both dendritic branches (see Figs. 3A, 4A and 4D for examples). Additionally, 108 feature-unspecific synapses were distributed throughout the dendrites, activated concurrently with all stimuli to enhance the probability of spiking (as plateau potentials together with the general background synaptic noise used do not often lead to spikes in SPNs, cf Figure 2A). The features representing a stimulus (clustered synapses) were active within 20 ms, while feature-unspecific input was activated over a 50 ms window. In simulations with inhibition, four inhibitory synapses, each one representing a feature, were placed near each cluster (Figure 5A). Within this setup a single stimulus activated both excitatory and inhibitory synapses. To match the level of depolarization seen in our excitatory-only setup, the number of feature-unspecific synapses were increased to 144 in the simulations including inhibitory plasticity. The initial conductance of inhibitory synapses was set to 0.1 ± 0.01 nS.

Revised version: This setup was based on the assumption of pre-existing clustered synapses for each feature. Features were allocated to two dendritic branches. Each branch had each feature represented with five synapses clustered closely on a single dendrite. Depending on the feature combination, two, three, or four features were represented in clusters on one or both dendritic branches (see Figure 3A, Figure 4A, and Figure 4D for examples). Additionally, 108 feature-unspecific synapses were distributed throughout the dendrites, activated concurrently with all stimuli to enhance the probability of spiking (as plateau potentials together with the general background synaptic noise used do not often lead to spikes in SPNs, cf Figure 2A). The features representing a stimulus (clustered synapses) were active within 20 ms, while feature-unspecific input was activated over a 50 ms window. In simulations with inhibition, four inhibitory synapses, each one representing a feature, were placed near each cluster (Figure 5A). Within this setup a single stimulus activated both excitatory and inhibitory synapses. To match the level of depolarization seen in our excitatory-only setup, the number of feature-unspecific synapses were increased to 144 in the simulations including inhibitory plasticity. The initial conductance of inhibitory synapses was set to 0.1 ± 0.01 nS.

Distributed setup

Earlier version: In contrast, this setup examined learning dynamics in neurons without pre-existing synaptic clustering for individual features. A total of 200 excitatory synapses were randomly distributed over 30 dendrites. Each feature was represented by 40 excitatory synapses, and an additional 40 feature-unspecific excitatory synapses were used (Figure 6A). We initiated this experiment with excitatory synaptic weights at 0.3 ± 0.1 (around 0.75 nS). This higher initial weight was chosen to compensate for the reduced efficacy of non-clustered synaptic inputs in producing sufficient depolarization and calcium influx. Extending our investigation, we added 60 inhibitory synapses, 15 for each feature, dispersed randomly over the 30 dendrites. We initiated this extended experiment with excitatory synaptic weights at 0.45 ± 0.1 (around 1.125 nS), aiming to maintain the same baseline voltage activity as in the excitatory-only case. Figure 6A illustrates this setup, exemplifying the pre- and post-learning synaptic weights for both excitatory and inhibitory synapses.

Revised version: In contrast, this setup examined learning dynamics in neurons without pre-existing synaptic clustering for individual features. A total of 200 excitatory synapses were randomly distributed over 30 dendrites. Each feature was represented by 40 excitatory synapses, and an additional 40 featureunspecific excitatory synapses were used (Figure 6A). We initiated this experiment with excitatory synaptic weights at 0.3 ± 0.1 (around 0.75 nS). This higher initial weight was chosen to compensate for the reduced efficacy of non-clustered synaptic inputs in producing sufficient depolarization and calcium influx. Extending our investigation, we added 60 inhibitory synapses, 15 for each feature, dispersed randomly over the 30 dendrites. We initiated this extended experiment with excitatory synaptic weights at 0.45 ± 0.1 (around 1.125 nS), aiming to maintain the same baseline voltage activity as in the excitatory-only case. Figure 6A illustrates this setup, exemplifying the pre- and post-learning synaptic weights for both excitatory and inhibitory synapses.

Added in the revised version: Code and Data Availability

Added in the revised version: The modeling code used in this study is structured for reproducibility and is available at https://github.com/zahradd/dSPN-learning-rule. This repository includes the simulations underlying the figures in the manuscript, along with a README file providing detailed instructions for environment setup and result reproduction. Additionally, model components adapted from Lindroos and Hellgren Kotaleski (2021) and Trpevski et al. (2023) are available on ModelDB (accession numbers: 266775 and 2017143).